# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">PHYS 121 Lab #3</span><span style="color:red"> $\to$ (14 possible marks)</span></div>

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# Electric Potential & Electric Fields

***
## Learning Objectives:
* <b><span style="color:rgba(0, 153, 76, 1);">Extract information about a physical system from a contour plot of a scalar field.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Determine the electric field in a region of space from a map of the electric potential.</span></b>

In this lab you will create maps of electric potentials on 2-D surfaces generated by various patterns of electrodes and conductors.  From these maps, you will identify lines of constant potential (i.e. equipotential lines) and deduce the configuration of the corresponding electric field.  Your measurements will be made using a sheet of conducting carbon paper, a DC power supply, and a handheld multimeter.

***
## Autograding:
The PHYS 121 Pre-lab assignments and Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS121.Installer()' reports that some functions have been installed, the user should restart the "kernel" and then exectute the PHYS121.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS121.Installer()' command requires the file 'PHYS121.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Import PHYS121.py and then run the installer function.
import PHYS121
PHYS121.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 121 - Lab 3.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
import pandas as pd
import scipy.stats as sp
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Name & Student Number</span></div>  

To ensure that both you and your partner receive credit for you lab work, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number as strings (i.e. place your names and student numbers between single or double quotes.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
PHYS121.graderCheck([last_name_1, first_name_1, student_number_1, last_name_2, first_name_2, student_number_2], ['last_name_1', 'first_name_1', 'student_number_1', 'last_name_2', 'first_name_2', 'student_number_2'], grader.check('name_and_student_number'))
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - The Electric Field as a Gradient Field (10 minutes)</span></div>



For the pre-lab, you were introduced to scalar fields and vectors fields by thinking about the terrain around a rock quarry. Recall that if we define the altitude above sea level at a position $(x,y)$ by the function $h(x,y)$, the vector field
<br>

\begin{align}
\mathbf{s}(x,y) = \frac{dh}{dx}\,\hat\imath + \frac{dh}{dy}\,\hat\jmath,
\end{align}
<br>
points in the direction of steepest ascent from the point $(x,y)$ and has a magnitude $\left|\mathbf{s}(x,y)\right|$ equal to the slope at that point along the direction of steepest accent. 

Now, suppose you stand on a hill and release a ball from rest. It should spontaneously begin to roll in the direction of steepest *descent* - that is, it will initially roll in the direction of the vector field $-\mathbf{s}(x,y)$.

Analogously, a positive charge released from rest will move in the direction that the electric potential *decreases* the most rapidly. This is exactly the direction the electric field $\mathbf{E}$ points since, for a two-dimensional electric potential $V(x,y)$
<br>

\begin{align}
\mathbf{E}(x,y) = -\left( \frac{dV}{dx}\,\hat\imath + \frac{dV}{dy}\,\hat\jmath \right)
\end{align}
<br>

That is, when we look at our electrical field maps, one way to interpret them is to think about what would happen to a *positive* charge if released from rest in the field. We can also see this in the Lorentz force law, where the force on a charged particle due to an electric field $\mathbf{E}$ (and in the absence of a magnetic field) is given by
<br>

\begin{align}
\mathbf{F}(x,y) = q\mathbf{E}(x,y) = -q\left( \frac{dV}{dx}\,\hat\imath + \frac{dV}{dy}\,\hat\jmath \right)
\end{align}
<br>

Notice that this equation also tells us that a *negative* charge ($q < 0$) will move in the opposite direction that a positive charge would move.

In the following three questions, you'll get a bit of practice thinking about the electric field around a variety of systems of point charges. If you like, you can map the electric field using this [simulation](https://phet.colorado.edu/en/simulations/charges-and-fields) to help you answer these questions, but you should try to use your imagination to picture the electrical field beforehand.

### <span style="color:purple">The 'PHYS121.graderCheck(...)' statements for **Questions 1.1** and **1.2**,  only check that answers have been entered in the correct format.  They do **NOT** check that you've entered the correct answer.</span> 


***
**<span style="color:blue">Question 1.1</span>** **<span style="color:red">(1 mark)</span>**  

The positive charge shown at the centre of the figure is fixed in place.  A second **positive** charge is released from rest at the position indicated by the **X**.

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-3/images/config1.png" width="750"/>

After it is released, the poitive charge placed at **X**:

(a) moves in a straight-line path away from the fixed charge. <br>
(b) moves in a straight-line path towards from the fixed charge. <br>
(c) moves away from the fixed charge, but follows a curved path.<br>
(d) moves towards the fixed charge, but follows a curved path.<br>
(e) follows a circular path around the fixed charge.

***
**<span style="color:blue">Answer 1.1:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, or e) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a1_1 = ...
```

```python
PHYS121.graderCheck([a1_1], ['a1_1'], grader.check('q1.1'))
```


***
**<span style="color:blue">Question 1.2</span>** **<span style="color:red">(1 mark)</span>**  

A **negative** charge is released from rest at the position indicated by the **X** (all other charges shown are fixed in place).

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-3/images/econfig2.png" width="550"/>

After it is released, the negative charge placed at **X**:

(a) moves in a straight-line path towards the collection of positive charges. <br>
(b) moves in a straight-line path towards the collection of negative charges. <br>
(c) moves towards the collection of positive charges, but follows a curved path.<br>
(d) moves towards the collection of negative charges, but follows a curved path.<br>
(e) follows an erratic/unpredictable path.

***
**<span style="color:blue">Answer 1.2:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, or e) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a1_2 = ...
```

```python
PHYS121.graderCheck([a1_2], ['a1_2'], grader.check('q1.2'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Electric Potential & Electric Field Map (60 minutes)</span></div>

In this part of the lab, you will create maps of the electric potential and electric field on a conductive 2-D surface.  The main pieces of equipment that you will work with are:
- Sheets of carbon paper.  Various patterns of conducting electrodes have been added to the carbon paper using a silver paint.
- A DC power supply used to apply fixed voltages (or potentials) to the electrodes.
- A digital multimeter used to probe the value of the electric potential at various points on the carbon paper.
- Various wires.

## Procedure:
- Select one of the seven available carbon paper graphs.
- Using the DC power supply, connect $10\rm\ V$ to the red wire that is anchored to your graph and $0\rm\ V$ (the black output of the power supply labeled '$-$') to the black wire. 
- Set the multimeter to 'DC Voltage' mode and use the free of end of the red banana cable connected to measure the voltage at different points on the carbon paper.  Avoid placing the end of the banana cable directly on the grid lines of the carbon paper.  The grid lines are not good conductors and lead to unreliable measurements.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-3/images/Fluke115.jpg" width="200"/>
<br>
<b>Fig. 1: A Fluke 115 handheld multimeter set to 'DC Volatge' mode.</b></center>
</p>

- You should have plenty of time to complete this lab, so collect lots of data to produce maps that are as detailed and accurate as possible.  The sheets of carbon paper contain grid lines that define the $x$ and $y$ positions on the paper in units of centimeters.  If you were to collect data at $2\rm\ cm$ intervals (recommended), you'd end up with approximately 144 data points for each graph.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-3/images/carbon_paper.png" width="600"/>
<br>
<b>Fig. 2: Photograph of one of the carbon-paper sheets mounted to a plastic board.  The grid on the sheet spans $0$ to $25\rm\ cm$ in the $x$-direction and $0$ to $20\rm\ cm$ in the $y$-direction.</b></center>
</p>

- Collect the data.  The $x$ and $y$ positions that you enter must be in units of **centimeters** and the voltages must be in units of **volts**.  
- Because we need to collect a lot of data points, it is easiest to first enter the data into a spreadsheet.  We will use Google Sheets.  These data can then be written to a csv (**c**omma-**s**eparated-**v**alues) file and then imported into the Jupyter notebook as a Pandas dataframe.  An Google Sheets template file called "Lab #3 mapping spreadsheet" has been created for you.  Here's a [link to the template](https://docs.google.com/spreadsheets/d/1IixIRP0p8YziNJ6h00YMnZZiSsF1xJUtjhxSrMJ6jLs/copy?usp=sharing).  You can use this file to enter your data.  It has been pre-populated with column names and some recommend $x$ and $y$ coordinate values.  Feel free to use the suggested coordinates, or to delete them and enter your own.  However, please leave the names at the top of the columns unchanged.

Here is a GIF showing how to save the data as a csv file after it has been entered into the Google Sheets document.  The data entered in the GIF is garbage data made up for the purposes of creating the GIF.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-3/gifs/csv.gif" width="1250"/>
</center>
</p>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.1</span>** **<span style="color:red">(1 mark)</span>**  

Exectute the cell below to import the data from the csv file into a Pandas dataframe.  Replace the ... with the name of the csv file between single or double quotes.  For example, if the csv file is called map1.csv, then the code in the cell below would look like:

```python
part2_filename = 'map1.csv'
part2_map = pd.read_csv(part2_filename)

# Display the dataframe
part2_map
```

*** Please do not change any of the variable names in the cell below ***

```python
part2_filename = ...
part2_map = pd.read_csv(part2_filename)

# Display the dataframe
# The first line below is used to ensure that the entire contents of the dataframe is displaced.
with pd.option_context('display.max_rows', None, 'display.max_columns', None):
    display(part2_map)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Generate the Electric Potential & Electric Field Map
To create the desired false-colour plot for the electric potential and the vector field plot for the electric field, the 'Mapping' function will be used.  Roughly speaking, the 'Mapping' function uses an interpolation to fill in the gaps in the electric potential data that you acquired.  The interpolated data is then used to
1. produce a false-colour plot with some contour lines
2. calculate and plot the corresponding electric field

For additional information about the 'Mapping' function, see **Appendix A** near the end of this notebook.

***
**<span style="color:blue">Question 2.2</span>** **<span style="color:red">(2 marks)</span>**  

Execute the cell below to generate the map of equipotential and electric field.

**In the cell below, delete the ... and replace it with the Graph Number of the carbon paper that you're using.**

```python
# Delete the ... in the line below and replace it with the Graph Number of the carbon paper that you're using.
graphNumber = ...

# Call a function that creates a filename that will be used to save the equipotential and electric field map that will be generated. 
part2_fig_name = PHYS121.extension([part2_filename], "png")[1]

# The line below assumes that the columns of the spreadsheet that you created are named 'x (cm)', 'y (cm)', and 'potential (V)'.
PHYS121.Mapping(part2_map["x (cm)"].to_numpy(), 
                part2_map["y (cm)"].to_numpy(), 
                part2_map["potential (V)"].to_numpy(), 
                graphNumber, 
                True,
                part2_fig_name
                )
```

<!-- END QUESTION -->

If there are holes in your electric potential map, or you want to add more detailed measurements in one area of the plot, you can go back and add more data to your DataFrame, either manually or by editing your .csv file. 

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.3</span>** **<span style="color:red">(1 mark)</span>**   

(a) In the electric potential/electric field map you created above, in which direction does the electric field point?  From high to low potential, or from low to high potential?

(b) What is the direction of the electric field relative to the lines of constant potential (equipotential lines)?

***
**<span style="color:blue">Answer 2.3:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Second Electric Potential & Electric Field Map (40 minutes)</span></div>

In **Part 3** of the lab, you will simply repeat the measurements of **Part 2** using one of the other carbon paper graphs.  Specifically:
- Select a carbon paper sheet with a graph number that is different than the one you used in **Part 2**.
- Measure the electric potential as a function of $x$ and $y$ position.
- Here, again, is the [link to the Google Sheets template](https://docs.google.com/spreadsheets/d/1IixIRP0p8YziNJ6h00YMnZZiSsF1xJUtjhxSrMJ6jLs/copy?usp=sharing).
- Use the 'Mapping' function to produce electric potential and electric field maps.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.1</span>** **<span style="color:red">(1 mark)</span>**  

Exectute the cell below to import the data from the csv file into a Pandas dataframe.  Replace the ... with the name of the csv file between single or double quotes.  For example, if the csv file is called map2.csv, then the code in the cell below would look like:

```python
part3_filename = 'map2.csv'
part3_map = pd.read_csv(part3_filename)

# Display the dataframe
part3_map
```

*** Please do not change any of the variable names in the cell below ***

```python
part3_filename = ...
part3_map = pd.read_csv(part3_filename)

# Display the dataframe
# The first line below is used to ensure that the entire contents of the dataframe is displaced.
with pd.option_context('display.max_rows', None, 'display.max_columns', None):
    display(part3_map)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.2</span>** **<span style="color:red">(2 marks)</span>**  

Execute the cell below to generate the map of equipotential and electric field.

**In the cell below, delete the ... and replace it with the Graph Number of the carbon paper that you're using.**

```python
# Delete the ... in the line below and replace it with the Graph Number of the carbon paper that you're using.
graphNumber = ...

# Call a function that creates a filename that will be used to save the equipotential and electric field map that will be generated. 
part3_fig_name = PHYS121.extension([part3_filename], "png")[1]

# The line below assumes that the columns of the spreadsheet that you created are named 'x (cm)', 'y (cm)', and 'potential (V)'.
PHYS121.Mapping(part3_map["x (cm)"].to_numpy(), 
                part3_map["y (cm)"].to_numpy(), 
                part3_map["potential (V)"].to_numpy(), 
                graphNumber, 
                True,
                part3_fig_name
                )
```

<!-- END QUESTION -->

If there are holes in your electric potential map, or you want to add more detailed measurements in one area of the plot, you can go back and add more data to your DataFrame, either manually or by editing your .csv file. 

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - Reflection and statement of contributions (20 minutes)</span></div>

<!-- BEGIN QUESTION -->

Questions 4.1 to 4.3  are used to reflect on the learning objectives stated at beginning of the notebook.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/93f29aaaec97a4eca7c20253be8c3f8510228118/Lab-3/images/chargedConductors.svg" alt="image info" style="width: 65%; height: auto;" /><br>
<b>Fig. 3: The map of the electric potential in the region surrounding a pair of charged conducting spheres.  Figure adapted from <a href = "https://commons.wikimedia.org/wiki/File:VFPt_metal_balls_largesmall_potential%2Bcontour.svg">Wikimedia Commons</a></b></center>
</p>


***
**<span style="color:blue">Question 4.1</span>** **<span style="color:red">(1 mark)</span>**    

Is the electric field stronger at point $P_1$ or point $P_2$ in Fig. 3?  Explain your reasoning using a couple of short sentences.

***
**<span style="color:blue">Answer 4.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 4.2</span>** **<span style="color:red">(1 mark)</span>**    

What is the direction of the electric field at point $P_3$ in Fig. 3?  Explain your reasoning using a couple of short sentences.

***
**<span style="color:blue">Answer 4.2:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/93f29aaaec97a4eca7c20253be8c3f8510228118/Lab-3/images/charged-rod.svg" alt="image info" style="width: 65%; height: auto;" /><br>
<b>Fig. 4: The electric field of a uniformly-charged rod.  Figure taken from <a href = "https://commons.wikimedia.org/wiki/File:VFPt_charged-rod.svg">Wikimedia Commons</a>.
</p>


***
**<span style="color:blue">Question 4.3</span>**  **<span style="color:red">(1 mark)</span>**      

In Fig. 4, does the electric potential change most rapidly (i) near the centre of the rod, (ii) near the ends of the rod, or (iii) at a point very far from the rod.  Explain your reasoning using a couple of short sentences.

***
**<span style="color:blue">Answer 4.3:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***
Below, you will briefly summarize the contributions that you and your partner made when completely this lab.  Please ensure that both partners have a chance to engage in all types of lab activities (developing the experimental methods/data collection strategies, hands-on interaction with the apparatus/data collection, and data processing/analysis in the Jupyter notebook).  Note that statements of contributions are common in science.  The figure below shows one example taken from a 2021 paper published in the journal *Nature Physics*.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/contributions.jpg" alt="image info" style="width: 60%; height: auto;" /><br>
<b>Fig. 5: A sample statement of author contributions from a recent journal article.</b></center>
</p>


<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 4.4(a):</span>**  

Answer this question only if you worked with a partner to complete this lab.

In two or three sentences describe the specific contributions of **Partner #1** to this lab.  For example, worked with the experimental apparatus to collect data for Part 1 of the lab, entered and processed the data in the Jupyter notebook for Part 2 of the lab, ...  

***
**<span style="color:blue">Answer 4.4(a):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  Please start be entering the name and student number of **Partner #1**.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 4.4(b):</span>**  

Answer this question only if you worked with a partner to complete this lab.

In two or three sentences describe the specific contributions of **Partner #2** to this lab.  

***
**<span style="color:blue">Answer 4.4(b):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  Please start be entering the name and student number of **Partner #2**.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 5 - Student Surveys (25 minutes)</span></div>

The last thing that we'd like to ask you to do as part of Lab #3 is complete the following two online surveys: 

### Survey #1 - Experience in the Lab

The first survey is used to collect feedback about the lab activities and the Jupyter notebook-based lab manuals.  The lab activities and method of delivery are both new, so we really do value your feedback.

Here is the link to [Survey #1 - Experience in the Lab](https://ubc.ca1.qualtrics.com/jfe/form/SV_ahLqXpi0lfic1Om)


### Survey #2 - Quality of your TA

The second survey is used to collect feedback about the quality of your TA.  The anonymous feedback will be shared with your TA so that they can keep doing the things that you like while simultaneously trying to improve any weaknesses. 

Here is the link to [Survey #2 - TA Quality](https://ubc.ca1.qualtrics.com/jfe/form/SV_1C6F1YitYpCkAHY)

Both of these surveys are also available in the **PHYS 121 Labs 2022W2** [Canvas](https://canvas.ubc.ca/) shell. If your partner is using the lab computer to complete the surveys, feel free to log into Canvas to complete the two surveys.


***
**<span style="color:blue">Question 5.1</span>** **<span style="color:red">(1 mark)</span>**

Partner #1 and Partneer #2 (if applicable), completed both of the online surveys.

(a) True  
(b) False

***
**<span style="color:blue">Answer 5.1:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, T, or F) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a5_1 = ...
```

```python
PHYS121.graderCheck([a5_1], ['a5_1'], grader.check('q5.1'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 6 - Submission (5 minutes)</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 6.1:</span>**  

We welcome your feedback on the PHYS 121 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 121 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 6.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Canvas gradebook.
- **Do NOT change the number of the .zip file.**
- **Do NOT modify the contents of the .zip file.**

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell (place your cursor in the cell and then hit Shift + Enter) to export 
# your submission.  Uploaded the .zip file that is generated to the gradebook of the Canvas shell for the PHYS 121 lab.
grader.export(files = [part2_filename, part3_filename, part2_fig_name, part3_fig_name, 'PHYS121_DataLogger.txt'])
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 7 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.  
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.  
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Mapping Function...</div>

A function for producing false-colour plots of electric potential and then calculating and displaying the corresponding electric field map.  The $x$ and $y$ position data must be entered in units of centimeters.  The $x$ data should span $0$ to $25\rm\ cm$ and the $y$ data should span $0$ to $20\rm\ cm$.  The electric potential data should be in units of volts and span the range from $0$ to $10\rm\ V$.  The function is called as follows:
```python
Mapping(x, y, V, graphNum = 0, vectorField = True, fig_file_name = 'figure.png')
```
The 'x', 'y', and 'V' inputs are lists or arrays of floats and are all required.  'graphNum' is an optional integer input.  If an integer from 1 to 7 is entered, the approximate positions of the silver-paint electrodes will be shown on the final plot.  If other optional input 'vectorField' is equal to True, then the electric field plot will be shown.  The function returns the a single output (the formatted figure):
```python
fig
```


### Mapping Example Implementation
The code block below shows an implementation of 'Mapping'.
```python
import numpy as np
import PHYS121

map1 = pd.read_csv("Graph1Data.csv")

PHYS121.Mapping(map1["x (cm)"].to_numpy(), map1["y (cm)"].to_numpy(), map1["potential (V)"].to_numpy(), 1, True, 'fig1.png')
```

If you're interested in generating your own contour plots with customized formatting, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Surface%20Plots.html.
***

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: February 6, 2024

Tests for name_and_student_number (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...


Tests for q1.1 (points: None)
```
>>> isinstance(a1_1, str)
True
```
failure_message: Your answer must be a string, i.e. surrounded by single or double quotes.
```
>>> test = True
>>> if isinstance(a1_1, str):
...     test = a1_1.lower() in ["a", "b", "c", "d", "e"]
>>> test
True
```
failure_message: Your answer must be a string of one of 'a', 'b', 'c', 'd', or 'e'.


Tests for q1.2 (points: None)
```
>>> isinstance(a1_2, str)
True
```
failure_message: Your answer must be a string, i.e. surrounded by single or double quotes.
```
>>> test = True
>>> if isinstance(a1_2, str):
...     test = a1_2.lower() in ["a", "b", "c", "d", "e"]
>>> test
True
```
failure_message: Your answer must be a string of one of 'a', 'b', 'c', 'd', or 'e'.


Tests for q5.1 (points: None)
```
>>> isinstance(a5_1, str)
True
```
failure_message: Your answer must be a string, i.e. between single or double quotes.
```
>>> test = True
>>> if isinstance(a5_1, str):
...     test = a5_1.lower() in ["a", "b", "t", "f"]
>>> test
True
```
failure_message: Your answer must be a string of one of 'a', 'b', 'T', or 'F'.
