# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">**PHYS 121 Lab #2**<span style="color:red"> $\to$ (22 possible marks)</span> </span></div> 

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# Oscillations of a Simple Pendulum &ndash; Week 2

***
## Learning Objectives:
* <b><span style="color:rgba(0, 153, 76, 1);"> Develop a method to collect a set of precision measurements.  Implement the method and, if necessary, make revisions and repeat the measurements. </span></b>
* <b><span style="color:rgba(0, 153, 76, 1);"> Collect a set of precision measurements that reveals the limitations of the theoretical model used to analyze the period of a simple pendulum. </span></b>

Today, you will have the opportunity to refine and iterate (if necessary) your measurements of the period of a pendulum to generate a relatively high-precision dataset that uncovers new physical insights that go beyond the material covered in first-year physics lectures.


***
## Autograding:
'PHYS121.Installer()' checks to see if some required Python packages/modules are already installed.  If they are, then the function outputs a message saying that its okay to proceed.  If one or more of the packages are missing, the function attempts to install them and then outputs a message that says the user should exectute the PHYS121.Installer() cell a second time.  The second time the installer function is run, it should simply confirm that the required packages are already installed.  

If necessary, the user can "restart the kernel" start with a fresh notebook in which no cells have been executed.  To restart the kernel, select **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

```python
# Import PHYS121.py and then run the installer function.
import PHYS121
PHYS121.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 121 - Lab 2.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
import pandas as pd
import scipy.stats as sp
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Name & Student Number</span></div>  

To ensure that both you and your partner receive credit for you lab work, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number as strings (i.e. place your names and student numbers between single or double quotes.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

```python
PHYS121.graderCheck([last_name_1, first_name_1, student_number_1], ['last_name_1', 'first_name_1', 'student_number_1'], grader.check('name_and_student_number_1'))
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
PHYS121.graderCheck([last_name_2, first_name_2, student_number_2], ['last_name_2', 'first_name_2', 'student_number_2'], grader.check('name_and_student_number_2'))
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Introduction</span></div>  
## Recap of important definitions:

**mean** $\mu$ 

* A measure of the *average* value of a set of data. If you take $N$ measurements $\nu_1, \nu_2, \dots, \nu_N$, the mean is defined by

$$ \mu = \frac{1}{N}\sum_{i=1}^N \nu_i = \frac{\nu_1 + \nu_2 + \cdots + \nu_N}{N} $$

**standard deviation** $\sigma$  

* A measure of the amount of variation in a data set. A small value of $\sigma$ indicates that the data is clustered near the average value, and a large value of $\sigma$ indicates that the data is more spread out. 

$$ \sigma = \sqrt{\frac{1}{N - 1}\sum_{i=1}^N \left(\nu_i - \mu\right)^2} $$

**standard error in the mean** $\sigma_\mu$  

* A measure of how confident we are in our estimate of the mean. You can read $\mu \pm \sigma_\mu$ as "our measurements suggest that if we were to collect an identical set of measurements, there is a $68\%$ chance that the resulting mean will fall between $\mu - \sigma_\mu$ and $\mu + \sigma_\mu$." As you might expect, the confidence interval gets smaller when the distribution of the experimental data gets narrower (i.e. $\sigma$ is reduced) *and* when we take more measurements.

$$ \sigma_{\mu} = \frac{\sigma}{\sqrt{N}} $$

**normal distribution**

* Often called a *Gaussian distribution* or a *bell curve*. When one takes a large number of measurements, it is very often the case that the results are approximately normally distributed. Two important features to keep in mind are:

    - (a) normal distributions are *symmetric* about the mean 
    - (b) approximately $68\%$ of results fall within one standard deviation of the mean, $95\%$ within two, and $99.7\%$ within three. This is sometimes called the $68$-$95$-$99.7$ rule.

***
## Motivation for today's Lab:
Among the primary goals of physics (and in science in general) is the construction of a coherent picture of how the world works. In this pursuit, we wish to *build*, *test*, and *revise* models and theories. While *all* models are necessarily only approximations of reality, these approximations (and thereby our understanding of the world) *can* be improved.

One vital aspect of this work is to seek out and investigate cases where our models fail, but we must be *certain* that what we're seeing isn't skewed by imprecise measurements or other experimental flaws. This is one of the many reasons we may wish to increase the precision of our experiments. (Note that deciding whether an experimental result is significant is, in practice, a more subtle matter than simply calculating $\sigma$ - for those who are interested in learning more, [this](https://news.mit.edu/2012/explained-sigma-0209) brief article from MIT discusses some of the complications).

In Lab #1, you explored some of the factors that affect the quality of the data collected. Specifically, you should have seen that the **standard deviation** depends on the *quality* of the measurements, but not the *quantity*, while the **standard error in the mean** is sensitive to both. This week, you will continue to experiment with the simple pendulum (Fig. 1). Using the procedure you developed at the end of Lab #1 you will gather high-precision data in order to test the limits of the simple model of the pendulum.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-1/images/pendulum_schematic.png" alt="image info" style="width: 35%; height: auto;" /><br>
<b>Fig. 1: A simple pendulum suspended from the origin of a protractor.</b></center>
</p>

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Discussion and Further Refinement of Methodology (15 minutes)</span></div>

In **Part 4** of Lab #1, you were asked to refine your data-collection methods with the objective of reducing the standard deviation in your data. First, recount the method you used last lab and *why* you chose that method. Discuss this method first with your lab partner, and then with at least one other group. When you've done this, complete **Question 1.1** below.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.1</span>** **<span style="color:red">(2 marks)</span>**  

After you're done discussing, briefly explain below *how* you will be taking your measurements, and *why* you expect this to increase the precision of your measurements relative to the methods used at the beginning of Lab #1. This may be identical to the method you developed at the end of Lab #1, completely different, or somewhere in between.  This question is worth 2 marks, ("How": 1 mark and "Why: 1 mark), so provide enough detail that your TA can properly access your plan and why you expect it to be effective.

***
**<span style="color:blue">Answer 1.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***
In last week's lab, you began investigating the claim that the period of oscillation $T$ of a simple pendulum is given by 
\begin{align}
T = 2\pi\sqrt{\dfrac{\ell}{g}},
\label{eq:period} \tag{1}
\end{align}
where $\ell$ is the length of the string and $g=9.81\rm\ m/s^2$ is the gravitational acceleration.  Notice that, according to Eq. $(1)$, the period depends on neither the mass $m$ of the pendulum bob nor the amplitude of the oscillations. 

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Data Collection (90 minutes)</span></div>

**Part 2** will comprise the majority of the lab. You will measure the period of oscillation for initial amplitudes $\theta_0$ from $5^\circ$ to $45^\circ$ in intervals of $5^\circ$.

As you proceed, you may recognize new challenges and develop better strategies. You will be asked to reflect on these challenges and strategies later on, so it will be beneficial to take notes as you work. To this end, a text box has been provided below. This won't be marked - it's there to help you keep track of any changes you make to your procedure. 

Exectute the cell below to create an array of the amplitude values.

```python
# Array of amplitudes in degrees
amplitude = np.array([5, 10, 15, 20, 25, 30, 35, 40, 45])
```

We anticipate that, for each amplitude, you will repeat the measurements many times.  The cell below shows how you can quickly obtain the mean $\mu$, standard deviation $\sigma$ and standard error $\sigma_\mu=\sigma/\sqrt{N}$ from your collected data.  We can use NumPy to find the mean and standard deviation.  However, to find the standard error, we have to use the SciPy package.  This package was already imported at the beginning of this notebook.  The standard error is then found used:
```python
sp.sem(...)
```
where ... is replaced with an array of data.  Note that 'sem' stands from 'Standard Error in the Mean'.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.1</span>** **<span style="color:red">(9 marks)</span>**  

Enter the data you collected into the cell below.  Your data should be entered as arrays using the syntax:
```python
periodXX = np.array([...])
```
where the XX represents the oscillation amplitude in degrees (i.e. $05$, $10$, $15$, ..., $45$) and the ... between the square brackets should be replaced by your measured data with each entry separared by a comma.

As you can see, there is a **lot** of data to collect. It is often the case in experimental science the data collection is time consuming and repetive.  We don't apologize for it -- there is a potential payoff beyond just collecting the available marks!  If you do a careful and thorough job collecting high-quality data, we hope that by the end of this lab you will have uncovered some interesting physics that goes beyond what is typcially covered in a typical first-year physics class. 

We'll create a Pandas dataframe that can be used to visulaize the data as you collect it.  We'll add columns to this dataframe for each initial angle $\theta_0$ used. 

*** Please do not change any of the variable names in the cell below ***

```python
# Array of periods for a 5 degree ampltiude
period05 = ...

df = pd.DataFrame({'5 degrees': period05})

length05 = len(period05)
mu05 = np.mean(period05)
sig05 = np.std(period05)
sem05 = sp.sem(period05)
print('T_05 =', "{:.5f}".format(mu05), '+/-', "{:.5f}".format(sem05), 's') # the format syntax is used to set the number of decimal places displated. 
print('Length:', length05)

# Display the dataframe
df
```

```python
# Array of periods for a 10 degree ampltiude
period10 = ...

df['10 degrees'] = period10

length10 = len(period10)
mu10 = np.mean(period10)
sig10 = np.std(period10)
sem10 = sp.sem(period10)
print('T_10 =', "{:.5f}".format(mu10), '+/-', "{:.5f}".format(sem10), 's') 
print('Length:', length10)

# Display the dataframe
df
```

```python
# Array of periods for a 15 degree ampltiude
period15 = ...

df['15 degrees'] = period15

length15 = len(period15)
mu15 = np.mean(period15)
sig15 = np.std(period15)
sem15 = sp.sem(period15)
print('T_15 =', "{:.5f}".format(mu15), '+/-', "{:.5f}".format(sem15), 's') 
print('Length:', length15)

# Display the dataframe
df
```

```python
# Array of periods for a 20 degree ampltiude
period20 = ...

df['20 degrees'] = period20

length20 = len(period20)
mu20 = np.mean(period20)
sig20 = np.std(period20)
sem20 = sp.sem(period20)
print('T_20 =', "{:.5f}".format(mu20), '+/-', "{:.5f}".format(sem20), 's') 
print('Length:', length20)

# Display the dataframe
df
```

```python
# Array of periods for a 25 degree ampltiude
period25 = ...

df['25 degrees'] = period25

length25 = len(period25)
mu25 = np.mean(period25)
sig25 = np.std(period25)
sem25 = sp.sem(period25)
print('T_25 =', "{:.5f}".format(mu25), '+/-', "{:.5f}".format(sem25), 's') 
print('Length:', length25)

# Display the dataframe
df
```

```python
# Array of periods for a 30 degree ampltiude
period30 = ...

df['30 degrees'] = period30

length30 = len(period30)
mu30 = np.mean(period30)
sig30 = np.std(period30)
sem30 = sp.sem(period30)
print('T_30 =', "{:.5f}".format(mu30), '+/-', "{:.5f}".format(sem30), 's') 
print('Length:', length30)

# Display the dataframe
df
```

```python
# Array of periods for a 35 degree ampltiude
period35 = ...

df['35 degrees'] = period35

length35 = len(period35)
mu35 = np.mean(period35)
sig35 = np.std(period35)
sem35 = sp.sem(period35)
print('T_35 =', "{:.5f}".format(mu35), '+/-', "{:.5f}".format(sem35), 's') 
print('Length:', length35)

# Display the dataframe
df
```

```python
# Array of periods for a 40 degree ampltiude
period40 = ...

df['40 degrees'] = period40

length40 = len(period40)
mu40 = np.mean(period40)
sig40 = np.std(period40)
sem40 = sp.sem(period40)
print('T_40 =', "{:.5f}".format(mu40), '+/-', "{:.5f}".format(sem40), 's') 
print('Length:', length40)

# Display the dataframe
df
```

```python
# Array of periods for a 45 degree ampltiude
period45 = ...

df['45 degrees'] = period45

length45 = len(period45)
mu45 = np.mean(period45)
sig45 = np.std(period45)
sem45 = sp.sem(period45)
print('T_45 =', "{:.5f}".format(mu45), '+/-', "{:.5f}".format(sem45), 's') 
print('Length:', length45)

# Display the dataframe
df
```

<!-- END QUESTION -->

***
**<span style="color:blue">Notes (optional):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell. 

***
Now that you have determined the average (mean) and standard error (uncertainty in the mean) for each of the oscillation amplitudes, we can construct arrays of these values and use them to plot the results. Execute the cell below to create these arrays.

```python
period = np.array([mu05, mu10, mu15, mu20, mu25, mu30, mu35, mu40, mu45])
errorBars = np.array([sem05, sem10, sem15, sem20, sem25, sem30, sem35, sem40, sem45])
```

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.2</span>** **<span style="color:red">(5 marks)</span>**  

The cell below is used to generate a scatter plot (with error bars) of the data you collected.  **Question 2.2** will be partially graded on the quality of the data you collect.  

This plot is similar to the plot that you produced in **Part 1** of Lab #1, expect not you're (hopefully) working with a higher-quality dataset. For additional information about the 'Scatter' function, see **Appendix A** near the end of this notebook.

```python
fig = PHYS121.Scatter(amplitude, period, yErrors = errorBars, xlabel = 'Amplitude', ylabel = 'Period', xUnits = 'degrees', yUnits = 'seconds')
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.3</span>** **<span style="color:red">(1 mark)</span>** 

Once you've collected all of your data, comment on the scatter plot above. Is there a clear dependence of the period on amplitude?  Are your observation consistent with Eq. (1) for the period of a pendulum?  How do today's results compare to the observations you made in **Part 1** of Lab #1?  In that experiment, you likely concluded that the oscillation period was independent of the amplitude.  

***
**<span style="color:blue">Answer 2.3:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.4</span>** **<span style="color:red">(1 mark)</span>** 

What do you notice about the relative sizes of the error bars as the oscillation amplitude $\theta_0$ increases? Is there a pattern? Consider why this is so - did new challenges arise as you increased the amplitude? What other factors may have contributed to this effect?

***
**<span style="color:blue">Answer 2.4:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Analysis of Results (10 minutes)</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.1</span>** **<span style="color:red">(1 mark)</span>**   

You should see that the period *increases* as the initial angle increases, contrary to Eq. (1). In Part 4, you will be guided through an explanation of the main reason for this discrepancy, but there are many other possible factors. In the box below, suggest a few ideas. Be sure to justify your answers. There is no single correct answer, here - don't worry if you're not exactly correct, the point is to develop intuition by thinking actively about the physics of the system. 

Note: if your data does *not* show that the $T$ increases with $\theta_0$ and you don't have time to repeat the experiment, you may look at a neighbouring group's data when answering this question. If you do so, be sure to cite their names in your answer.


***
**<span style="color:blue">Answer 3.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - The Small Angle Approximation (25 minutes)</span></div>

****
<span style="color:rgba(220, 50, 50, 1);">**Note:**</span> this section contains some math which may be unfamiliar. Do not dispair! You won't be asked to calculate anything difficult.  We will only ask you one simple question! This section is meant to be an introduction to some of the new physics and math that can be used to understand some of your experimental discoveries!
****


Depending on how much math you've taken, you may heard of something called a "Taylor series." The basic idea is that most functions can be written as infinite polynomials:

$$ f(x) = a_0 + a_1\,x + a_2\,x^2 + \cdots + a_n\,x^n + \cdots = \sum_{n = 0}^{\infty}a_n\,x^n \tag{2} $$ 

For anyone interested in learning more, [this video](https://www.youtube.com/watch?v=3d6DsjIBzJ4) from [**3Blue1Brown**](https://www.youtube.com/c/3blue1brown?app=desktop) provides an intuitive introduction to Taylor series. 

Taylor series find many applications in physics. They frequently show up for the simple fact that polynomials are relatively simple to work with, so that **writing a function as a Taylor series can turn a difficult problem into an easy problem.**

If you're only considering values of $x$ which are very small (i.e. close to zero), you can usually get a good approximation by keeping only the first couple of terms in the Taylor series representation of $f(x)$.  This happens because a small number squared (or cubed) is much less than the original number. For example, we might write

$$ f(x) \approx a_0 + a_1\,x  \tag{3} $$

where we have ignored of the higher powers of $x$ (like $x^2$, $x^3$, ...).  This is sometimes referred to as a **first-order approximation**. The derivation of Eq. (1) resulted from just such an approximation. 

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-2/images/fbd.png" alt="image info" style="width: 35%; height: auto;" /><br>
<b>Fig. 2: Free-body diagram for the simple pendulum. </b></center>
</p>

From the free-body of Fig. 2, the force on the pendulum bob is given by 

$$ ma =  mg\sin\theta  \tag{3} $$

Because the pendulum bob is following a circular path, the acceleration can be expressed as $a = \ell\dfrac{d^2\theta}{dt^2}$, where $\ell$ is the length of the string.  Therefore, the motion of the pendulum is determined by:

$$ \frac{d^2\theta}{dt^2} =  \frac{g}{\ell}\sin\theta  \tag{4} $$

which is an expression that we derivied in class at the start of term.  It is *possible* to solve Eq. (4), but it isn't easy to do! Fortunately, the sine function can be expressed as a Taylor series as follows

$$ \sin \theta = \theta - \frac{\theta^3}{3!} + \frac{\theta^5}{5!} - \frac{\theta^7}{7!} + \cdots $$

Eq. (1) results from taking only the first-order term: this is often called the **small-angle approximation**:

$$ \sin\theta \approx \theta$$

Figure 3 shows a geometric way to achieve the same result:

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-2/images/saa2.png" alt="image info" style="width: 45%; height: auto;" /><br>
<b>Fig. 3: Geometric demonstration of the small angle approximation. Consider what happens as the angle $2\theta$ gets smaller: the length of the blue arc will approach the length of the red chord (in exactly the same way that the earth seems flat when you're standing on it). The length of the arc is $r\theta$, and the length of the chord is $r\sin\theta.$ Thus, for small $\theta$, $r\sin\theta \approx \theta.$ We can cancel the radius $r$ to obtain $\sin\theta \approx \theta.$ </b></center>
</p>

***   
Type shift+enter in the cell below to see $\theta$, $\theta - \theta^3/3!$, and $\sin\theta$ all plotted as a function of $\theta$ in a single graph.

```python
theta = np.arange(0, 45, 0.5)

order1 = np.radians(theta) # theta in radians
order2 = np.radians(theta) - (np.radians(theta)**3)/6 # theta - theta^2/6 in radians
exact = np.sin(np.radians(theta)) # sin(theta)

fig, ax = plt.subplots(figsize = (7, 7))
ax.plot(theta, exact, color = 'k', label = (r"$\sin\theta$"), linewidth = 3)
ax.plot(theta, order1, 'r--', label = (r"$\theta$"), linewidth = 3)
ax.plot(theta, order2, 'c:', label = (r"$\theta - \theta^3/6$"), linewidth = 5)
ax.legend()

plt.xlim([0, 45])
plt.ylim([0, 0.75])
plt.xlabel(r"$\theta$")
plt.show()
```

<!-- BEGIN QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 4.1</span>**  **<span style="color:red">(1 mark)</span>**  

(a) Based on the plot above, over what range of angles (in degrees) is the approximation $\sin\theta = \theta$ valid?  

(b) Since $ T = 2\pi(\ell/g)^{1/2} $ was derived assuming $\sin\theta\approx\theta$, over what range of angles (in degrees) would you expect this expression to be valid?

***
**<span style="color:blue">Answer 4.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***
For anyone interested, the *exact* solution of Eq. (4) gives a period of 

$$ T = 2\pi\sqrt{\frac{\ell}{g}}\left(1 + \frac{1}{16}\theta_0^2 + \frac{11}{3027}\theta_0^4 + \frac{173}{737280}\theta_0^6 +\cdots \right) $$

where $\theta$ is in radians.

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 5 - Reflection & statement of contributions (25 minutes)</span></div>

<!-- BEGIN QUESTION -->

Question 5.1 is used to reflect on the learning objectives stated at beginning of the notebook. 
***
**<span style="color:blue">Question 5.1</span>** **<span style="color:red">(1 mark)</span>**   

Compare the scatter plot of period versus amplitude that you produced in **Part 1** of Lab #1 to the plot produced in **Part 2** of this lab. In a few short sentences, describe how increasing the precision of your measurements led to the discovery of new physics and revealed a limitation of the model often used to analyze simple pendulums.

***
**<span style="color:blue">Answer 5.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***
Below, you will briefly summarize the contributions that you and your partner made when completely this lab.  Please ensure that both partners have a chance to engage in all types of lab activities (developing the experimental methods/data collection strategies, hands-on interaction with the apparatus/data collection, and data processing/analysis in the Jupyter notebook).  Note that statements of contributions are common in science.  The figure below shows one example taken from a 2021 paper published in the journal *Nature Physics*.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/contributions.jpg" alt="image info" style="width: 60%; height: auto;" /><br>
<b>Fig. 2: A sample statement of author contributions from a recent journal article.</b></center>
</p>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 5.2(a):</span>**  

Answer this question only if you worked with a partner to complete this lab.

In two or three sentences describe the specific contributions of **Partner A** to this lab.  For example, worked with the experimental apparatus to collect data for Part 1 of the lab, entered and processed the data in the Jupyter notebook for Part 2 of the lab, ...  

***
**<span style="color:blue">Answer 5.2(a):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  Please start be entering the name and student number of **Partner A**.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 5.2(b):</span>**  

Answer this question only if you worked with a partner to complete this lab.

In two or three sentences describe the specific contributions of **Partner B** to this lab.  

***
**<span style="color:blue">Answer 5.2(b):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  Please start be entering the name and student number of **Partner B**.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 5.3:</span>**  

We welcome your feedback on the PHYS 121 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 121 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 5.3:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Canvas gradebook.
- **Do NOT change the number of the .zip file.**
- **Do NOT modify the contents of the .zip file.**

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(files = ['PHYS121_DataLogger.txt'])
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 6 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.  
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.  
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS121
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS121.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: February 1, 2024

Tests for name_and_student_number_1 (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)


Tests for name_and_student_number_2 (points: None)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...
