# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">**PHYS 121 Lab #0**<span style="color:red"> $\to$ (15 possible marks)</span> </span></div> 

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-0/gifs/PHYS121.gif" alt="image info" style="width: 85%; height: auto;" />
</center>
</p>

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

***
## Lab Learning Outcomes:
By the end of the PHYS 121 lab, the successful student will be able to:
* <b><span style="color:rgba(0, 153, 76, 1);">Take familiar concepts and techniques used to solve a problem in one field physics and use them to gain insights about problems from other, seemingly unrelated, fields of physics.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Use basic/inexpensive equipment to conduct careful physics experiments that produce meaningful results that uncover something about the behaviour of nature.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Revise and repeat experiments using improved measurement techniques to collect precision data that uncovers new insights about the physical system being studied.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Use Python and Jupyter Notebooks to analyze the experimental data that you collect.</span></b>

## Lab Manuals:
The PHYS 121 lab manuals have been created using interactive **Jupyter Notebooks** (https://en.wikipedia.org/wiki/Project_Jupyter).  These notebooks, which have an **.ipynb** extension, allow one to enter, analyze/process, and display experimental data.  You will be provided with links that will open the electronic notebook/lab manual.  You will enter your data and do your analysis within these notebooks.  At the end of each lab, you will upload a **.zip** file to the PHYS 121 Lab Canvas page.  The .zip file will be generated for you, and it will contain the .ipynb file that you filled in with your data and anaylsis and a pdf file with some of the results from your analysis. 

Jupter notebooks have "cells" that are used to enter text, do calculations, and/or analyze data.  To execute a Juyter cell, place your cursor within the cell and then hit SHIFT + ENTER (hold down SHIFT and then press ENTER).  Each cell can contain more than one line of text/code.  To start a new line, simply hit ENTER (without holding down the SHIFT key).

The Jupyter Notebooks that we will use in PHYS 121 will execute code written in the Python programming language (https://en.wikipedia.org/wiki/Python_(programming_language)).  You do **NOT** need to know Python or have any programming experience to successfully complete the PHYS 121 labs.  We will provide you with the code needed to complete the lab exercises.  Comments can be entered into a code cell by preceeding the comment with '#'.  Any text following the comment tag will not be executed as code when the user hits SHIFT + ENTER.

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Installing & Importing Packages (10 minutes)</span></div>

## The 'PHYS121.Installer()' Function:

**Part 1** of this introductory lab is used to install and import some packages that we will use throughout PHYS 121.  The installation will be done by calling a function called 'Installer()' which is contained within the file 'PHYS121.py'.  In order to use 'PHYS121.py', the file must be located in the same directory as the active Jupyter notebook (i.e. the .ipynb file) and then it must be imported into the notebook.  The 'PHYS121.py' file should already be among the files listed in the frame along the left-hand side of the screen.  The file can be imported using the following syntax:
```python
import PHYS121
```
The 'Installer()' function can then be called using:
```python
PHYS121.Installer()
```
These two line of code have already been entered into the cell below.  

'PHYS121.Installer()' checks to see if some required Python packages/modules are already installed.  If they are, then the function outputs a message saying that its okay to proceed.  If one or more of the packages are missing, the function attempts to install them and then outputs a message that says the user should exectute the PHYS121.Installer() cell a second time.  The second time the installer function is run, it should simply confirm that the required packages are already installed.  

If necessary, the user can "restart the kernel" start with a fresh notebook in which no cells have been executed.  To restart the kernel, select **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

Run 'PHYS121.Installer()' by placing you're cursor inside the code cell below and then hitting 'Shift'+'Enter'.

```python
import PHYS121
PHYS121.Installer()
```

## Importing Packages/Modules:

Once the packages are installed, they can be imported using the notation:
```python
import packageName
```
where 'packageName' is the name of the package to be imported.  In the cell below, we import some of the packages that were installed by 'PHYS121.Installer()' and others that are preinstalled within UBC's Open Jupyter.

Execute the import statements below by placing your cursor within the code cell and hitting 'Shift' + 'Enter'.

```python
# Import required and commonly-used modules
import numpy as np
import matplotlib.pyplot as plt
import uncertainties
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Auto-Grading (5 minutes)</span></div>

The PHYS 121 Pre-Lab Assignments and Lab Notebooks will make use of both auto-graded questions and manually-graded questions.  In order to facilitate the auto-grading and generate the files for manual grading, the 'otter' package has to be imported and then initiallized using the commands:
```python
import otter
grader = otter.Notebook("PHYS 121 - Lab 0.ipynb")
```
Execute the cell below (SHIFT + ENTER) to import and initialize the auot-grading. 

```python
# Import Otter for the purposes of autograding
import otter

# Initialize Otter (for the purposes of auto-grading some cells)
grader = otter.Notebook("PHYS 121 - Lab 0.ipynb")
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Lab Logistics (10 minutes)</span></div>  

In the PHYS 121 lab, most of you will be working with a partner.  **If you are working with a partner, you only need to upload a single copy of the zip file to Canvas when submitting your work.**  The zip file can be uploaded to either your Canvas shell or your partner's Canvas shell.  Once the submission has been graded, the same mark will be entered for both you and your partner.  Note that you need to complete the **pre-lab assignments independently**.  The pre-labs will also be completed using Jupyter notebooks.  For the pre-labs, every student in the PHYS 121 lab must upload their own zip file to Canvas.

To ensure that both you and your partner receive credit for you lab work, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number.

After your names and student numbers have been entered, execute the cells containing
```python
grader.check("name_and_student_number_1")
```
and
```python
grader.check("name_and_student_number_2")
```
These fucntions are part of the auto-grader and are desiged to check for errors in the data that you entered.  For example, the student number should be a string (enclosed in single or double quotes) of exactly 8 characters and all of the characters have to be integers.  If the data that you entered is okay, you should get a message indicating that all of the checks have passed.  If anything fails, you should search for the **message** in the output and use it to fix the error(s) in the values that you entered.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

```python
grader.check("name_and_student_number_1")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
grader.check("name_and_student_number_2")
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 Quiz (45 minutes)</span></div>

In order for you to get familiar with the Jupyter Notebook lab manuals and the PHYS 121 policies, you will complete the short quiz below.

First, access the PHYS 121 Lab Canvas page.  Navigate the to "Lab 0" folder within the "Modules" section of Canvas and read the pdf documents there.  These documents contain the answers to most of the questions in the quiz.

To enter a response to the quiz questions, replace the dots (...) with the your answer (keep the answer within the singe quotes ' ').  Text inside of single quotes (' ') or double quotes (" ") are called strings.  The autograder expects your answers to be enter as strings.

***
**<span style="color:blue">Question 4.1</span>** **<span style="color:red">(1 mark)</span>**  
What is a passing grade for the PHYS 121 lab?
- (a) 35%
- (b) 40%
- (c) 50%
- (d) 60%
- (e) 65%

**<span style="color:blue">Answer 4.1</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, or e) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_1 = ...
```

```python
grader.check("Q4.1")
```

***
**<span style="color:blue">Question 4.2</span>** **<span style="color:red">(1 mark)</span>**  
There are no exemptions for the PHYS 121 lab.
- (a) True
- (b) False

**<span style="color:blue">Answer 4.2</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, t, or f) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_2 = ...
```

```python
grader.check("Q4.2")
```

***
**<span style="color:blue">Question 4.3</span>** **<span style="color:red">(1 mark)</span>**  
I am not required to attend every lab in person.
- (a) True
- (b) False

**<span style="color:blue">Answer 4.3</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, t or f) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_3 = ...
```

```python
grader.check("Q4.3")
```

***
**<span style="color:blue">Question 4.4</span>** **<span style="color:red">(1 mark)</span>**  
What is the due date and time for this lab?
- (a) 1 hour after the start of the lab.
- (b) 3 hours after the start of the lab.
- (c) 15 minutes after the end of the lab.
- (d) 23:59 the day before the next lab starts.
- (e) By the end of the term.

**<span style="color:blue">Answer 4.4</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, or e) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_4 = ...
```

```python
grader.check("Q4.4")
```

***
**<span style="color:blue">Question 4.5</span>** **<span style="color:red">(1 mark)</span>**  
To make special arrangements with your TA or the Physics Lab Team, I should:
- (a) make the arrangements verbally.
- (b) ask my friend.
- (c) have my mom or dad contact the university.
- (d) use the Qualtrics form ([Link to Form](https://ubc.ca1.qualtrics.com/jfe/form/SV_bNJfeDCS8TW0FwO)).
- (e) Special arrangements are not needed.  I can just show up to another lab section.

**<span style="color:blue">Answer 4.5</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, or e) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_5 = ...
```

```python
grader.check("Q4.5")
```

***
**<span style="color:blue">Question 4.6</span>** **<span style="color:red">(1 mark)</span>**  
For each calendar day that a lab is submitted late, there is late penalty deduction of:
- (a) 0%
- (b) 10%
- (c) 15%
- (d) 20%
- (e) 25%

**<span style="color:blue">Answer 4.6</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, or e) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_6 = ...
```

```python
grader.check("Q4.6")
```

***
**<span style="color:blue">Question 4.7</span>** **<span style="color:red">(1 mark)</span>**  
To be considered for a make-up lab, when you miss a lab you must declare your absence using the Qualtrics form ([Link to Form](https://ubc.ca1.qualtrics.com/jfe/form/SV_bNJfeDCS8TW0FwO)) within how many calendar days from the date of the lab that you missed?
- (a) 1 day
- (b) 2 days
- (c) 3 days
- (d) 4 days
- (e) 5 days

**<span style="color:blue">Answer 4.7</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, or e) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_7 = ...
```

```python
grader.check("Q4.7")
```

***
**<span style="color:blue">Question 4.8</span>** **<span style="color:red">(1 mark)</span>**  
An act of dishonesty (such as plagiarism) is not allowed. For this reason, it may result in a failing grade or mark of zero on the assignment or in the course. Cases of academic dishonesty will be reported to the Dean's office.
- (a) True
- (b) False

**<span style="color:blue">Answer 4.8</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, t, or f) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_8 = ...
```

```python
grader.check("Q4.8")
```

***
**<span style="color:blue">Question 4.9</span>** **<span style="color:red">(1 mark)</span>**  
The due date for finalizing lab grades is April 9. You must talk to your TAs to discuss any issues related to the lab marks before this date. We will not discuss it after the finalizing date.
- (a) True
- (b) False

**<span style="color:blue">Answer 4.9</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, t, or f) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_9 = ...
```

```python
grader.check("Q4.9")
```

***
**<span style="color:blue">Question 4.10</span>** **<span style="color:red">(1 mark)</span>**  
I am able to make a special arrangement at the end of term so that I can submit assignments that I did not hand earlier in the term to increase my grade just enough to pass the lab component of PHYS 121.
- (a) True
- (b) False

**<span style="color:blue">Answer 4.10</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, t, or f) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_10 = ...
```

```python
grader.check("Q4.10")
```

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 4.11</span>** **<span style="color:red">(1 mark)</span>**  
Enter the lab section that you are registered in.

**<span style="color:blue">Answer 4.11</span>**  
Replace the ... in the cell below with your answer.  Your answer should be a string of 3 characters ('L04', for example).  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_11 = ...
```

```python
grader.check("Q4.11")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 4.12</span>** **<span style="color:red">(1 mark)</span>**  
Enter your TAs first and last name.

**<span style="color:blue">Answer 4.12</span>**  
Replace the ... in the cell below with your answers.  Your answers should be strings (i.e. between single or double quotes).  
*** Please do not change anything to the left of the equals signs. ***

```python
TA_FIRST_name = ...
TA_LAST_name = ...
```

```python
grader.check("Q4.12")
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 5 - Mathematical Operations in Python (30 minutes)</span></div>

In Python, you can evaluate most basic mathematical operations in the way that you'd expect.  For example, try executing the cells below. 

```python
2 + 3
```

```python
2*3
```

```python
2/3
```

In many programming languages (and on Texas Instruments graphing calculators), a number is raised to a power using the notation $x\hat{}y$.  For example, $2\hat{}3$ would imply $2^3 = 8$.  In Python, however, a number is raised to a power by repeating the multiplication symbol as in:
```python
x**y
```
Confirm this behaviour by executing the cell below.

```python
2**3
```

Special function such as $\sin x$, $\cos x$, $\ln x$, $\exp(x)$ and so on will also not work as expected.  These functions, however, are defined within very nice modules or packages that have been written and are available for us to use.  One package that we'll almost always want to use is called **NumPy** (https://en.wikipedia.org/wiki/NumPy).  The constant $\pi$ is also defined within NumPy.  To use this package, we first need to import it.  NumPy is usually given the shorthand name 'np'.  To import the package and give it the name np, we need to execute the following line of code:
```python
import numpy as np
```
This has already been done in **Part 1** of this lab, so we don't need to repeat it here.  Execute the cells below to see how to use NumPy and special functions.

```python
np.pi
```

```python
np.cos(1)
```

```python
# If we want, we can put a number of statements on a single line.
np.sin(0), np.sin(1), np.sin(np.pi/2)
```

```python
np.exp(1)
```

Note that the in Python, and other programming languages as well, $\log x$ is used to denote the *natural logarithm* of $x$ or $\ln x$.  For example, executing the cell below should result in and answer of $2$.

```python
np.log( np.exp(2) )
```

To take the logaritm of a number using in the base-10 number system, one would use:
```python
np.log10(x)
```
as in the cell below.

```python
np.log10( 100 ), np.log10( 10**2 )
```

Let's practice basic mathematical operations by trying a **simple math trick**.
***

**<span style="color:blue">Question 5.1</span>** **<span style="color:red">(1 mark)</span>**  
Pick any 3-digit integer.  Write the number twice in a row and enter it into the cell below (eg. $372 \to 372372$).

**<span style="color:blue">Answer 5.1</span>**  
Replace the ... in the cell below with your answer.  For this answer, an integer should be entered and quotation marks are not needed.  
*** Please do not change anything to the left of the equals sign. ***

```python
x = ...
```

```python
grader.check("Q5.1")
```

***

**<span style="color:blue">Question 5.2</span>** **<span style="color:red">(1 mark)</span>**  
Take which ever number you picked from **Question 5.1** ($x$) and in the code cell below divide it first by 13, then divide by 11, and lastly divide by 7.

**<span style="color:blue">Answer 5.2</span>**  
Replace the ... with a mathematical expression that will, all at once, preform the entire operation described above.  
*** Please do not change anything to the left of the equals sign. ***

```python
y = ...

# Print the result of the calculation.
print(y)
```

```python
grader.check("Q5.2")
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 6 - Calling Functions Contained in PHYS121.py (20 minutes)</span></div>

**Part 6** is just for fun and won't be graded.  It is designed to show you how to call functions contained within the file PHYS121.py.  We will rely on many of the functions in PHYS121.py to help us analyze the experimental data that we collect through the term in the PHYS 121 lab.

First, execute the cell below.  It calls a function contained in PHYS121.py called 'printDigits()'.  In order to call this function (or any other function in the file), we first have to import PHYS121.py use:
```python
import PHYS121
```
This import was already done in **Part 1** of the lab, so we don't have to do it again.  The 'printDigits()' function generates a random list of integers between 1 and 9.  The list of integers will contain between 25 and 35 entries.  If, for some reason, you don't like the set of integers that was generated, you can simply keeping executing the cell until you get a set of numbers that you like.  'printDigits()' will also calculate the product of the set of integers that you generated.   

```python
PHYS121.printDigits()
```

***
Next, take the product from the previous step and change any **non-zero** digit in that product to zero.  For example, if your product was $340202565$, you might decide to change the $4$ to zero such that your product gets changed to $300202565$.

The next function that we'll make use of is called 'chase()'.  It is called in the same way as the previous function, except this time we need to pass an *arugument* or *parameter* to the function.  The argument is passed by entering your number as a string (i.e. between single or double quotes) between the brackets as in:
```python
num = '300202565'
PHYS121.chase(num)
```
Call the 'chase()' function and pass the modified product (the product after changing one non-zero digit to zero) as the argument.  'chase()' should be able to tell you which digit to secretly changed to zero.

Replace ... with your modified product formatted as a string.  
*** Please do not change anything to the left of the equals sign. ***

```python
num = ...
PHYS121.chase(num)
```

```python
grader.check("Q6.1")
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 7 - Statement of contributions (15 minutes)</span></div>

Below, you will briefly summarize the contributions that you and your partner made when completely this lab.  Please ensure that both partners have a chance to engage in all types of lab activities (developing the experimental methods/data collection strategies, hands-on interaction with the apparatus/data collection, and data processing/analysis in the Jupyter notebook).  Note that statements of contributions are common in science.  The figure below shows one example taken from a 2021 paper published in the journal *Nature Physics*.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/contributions.jpg" alt="image info" style="width: 60%; height: auto;" /><br>
<b>Figure: A sample statement of author contributions from a recent journal article.</b></center>
</p>


<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 7.1(a):</span>**  

Answer this question only if you worked with a partner to complete this lab.

In two or three sentences describe the specific contributions of **Partner #1** to this lab.  For example, worked with the experimental apparatus to collect data for Part 1 of the lab, entered and processed the data in the Jupyter notebook for Part 2 of the lab, ...  

***
**<span style="color:blue">Answer 7.1(a):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  Please start be entering the name and student number of **Partner #1**.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer.

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 7.1(b):</span>**  

Answer this question only if you worked with a partner to complete this lab.

In two or three sentences describe the specific contributions of **Partner #2** to this lab.  

***
**<span style="color:blue">Answer 7.1(b):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  Please start be entering the name and student number of **Partner #2**.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 8 - Feedback and Submission (15 minutes)</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 8.1:</span>**  

We welcome your feedback on the PHYS 121 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 121 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 8.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 


<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Canvas gradebook.
- **Do NOT change the number of the .zip file.**
- **Do NOT modify the contents of the .zip file.**

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 9 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way you like.  

```python
# Here's an empty code cell that you can use.
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: January 5, 2024

Tests for Q4.1 (points: None)
```
>>> a4_1 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_1, str)
True
```
failure_message: a1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_1, str):
...     test = a4_1.lower() in ['a', 'b', 'c', 'd', 'e']
>>> test
True
```
failure_message: a1 must be one of 'a', 'b', 'c', 'd', or 'e'.
```
>>> PHYS121.get_hash(a4_1.lower()) == '4a8a08f09d37b73795649038408b5f33'
True
```
failure_message: Try again.


Tests for Q4.10 (points: None)
```
>>> a4_10 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_10, str)
True
```
failure_message: a10 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_10, str):
...     test = a4_10.lower() in ['a', 'b', 't', 'f']
>>> test
True
```
failure_message: a10 must be one of 'a', 'b', 't', or 'f'.
```
>>> PHYS121.get_hash(a4_10.lower()) in ['92eb5ffee6ae2fec3ad71c777531578f', '8fa14cdd754f91cc6554c9e71929cce7']
True
```
failure_message: Try again.


Tests for Q4.11 (points: None)
```
>>> a4_11 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_11, str)
True
```
failure_message: a11 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_11, str):
...     test = len(a4_11) == 3
>>> test
True
```
failure_message: a11 must be a string of exactly 3 characters.
```
>>> test = True
>>> if isinstance(a4_11, str):
...     test = a4_11 in ['L01', 'L02', 'L03', 'L04', 'L05', 'L06', 'L07', 'L08', 'L09', 'L10']
>>> test
True
```
failure_message: a11 must be one of 'L01', 'L02', ..., 'L10'.


Tests for Q4.12 (points: None)
```
>>> TA_FIRST_name != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(TA_FIRST_name, str)
True
```
failure_message: TA_FIRST_name must be a string (i.e. surround your answer by either single or double quotes).
```
>>> TA_LAST_name != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(TA_LAST_name, str)
True
```
failure_message: TA_LAST_name must be a string (i.e. surround your answer by either single or double quotes).


Tests for Q4.2 (points: None)
```
>>> a4_2 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_2, str)
True
```
failure_message: a2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_2, str):
...     test = a4_2.lower() in ['a', 'b', 't', 'f']
>>> test
True
```
failure_message: a2 must be one of 'a', 'b', 't', or 'f'.
```
>>> PHYS121.get_hash(a4_2.lower()) in ['92eb5ffee6ae2fec3ad71c777531578f', '8fa14cdd754f91cc6554c9e71929cce7']
True
```
failure_message: Try again.


Tests for Q4.3 (points: None)
```
>>> a4_3 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_3, str)
True
```
failure_message: a3 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_3, str):
...     test = a4_3.lower() in ['a', 'b', 't', 'f']
>>> test
True
```
failure_message: a3 must be one of 'a', 'b', 't', or 'f'.
```
>>> PHYS121.get_hash(a4_3.lower()) in ['92eb5ffee6ae2fec3ad71c777531578f', '8fa14cdd754f91cc6554c9e71929cce7']
True
```
failure_message: Try again.


Tests for Q4.4 (points: None)
```
>>> a4_4 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_4, str)
True
```
failure_message: a4 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_4, str):
...     test = a4_4.lower() in ['a', 'b', 'c', 'd', 'e']
>>> test
True
```
failure_message: a4 must be one of 'a', 'b', 'c', 'd', or 'e'.
```
>>> PHYS121.get_hash(a4_4) == '4a8a08f09d37b73795649038408b5f33'
True
```
failure_message: Try again.


Tests for Q4.5 (points: None)
```
>>> a4_5 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_5, str)
True
```
failure_message: a5 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_5, str):
...     test = a4_5.lower() in ['a', 'b', 'c', 'd', 'e']
>>> test
True
```
failure_message: a5 must be one of 'a', 'b', 'c', 'd', or 'e'.
```
>>> PHYS121.get_hash(a4_5.lower()) == '8277e0910d750195b448797616e091ad'
True
```
failure_message: Try again.


Tests for Q4.6 (points: None)
```
>>> a4_6 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_6, str)
True
```
failure_message: a6 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_6, str):
...     test = a4_6.lower() in ['a', 'b', 'c', 'd', 'e']
>>> test
True
```
failure_message: a6 must be one of 'a', 'b', 'c', 'd', or 'e'.
```
>>> PHYS121.get_hash(a4_6.lower()) == 'e1671797c52e15f763380b45e841ec32'
True
```
failure_message: Try again.


Tests for Q4.7 (points: None)
```
>>> a4_7 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_7, str)
True
```
failure_message: a7 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_7, str):
...     test = a4_7.lower() in ['a', 'b', 'c', 'd', 'e']
>>> test
True
```
failure_message: a7 must be one of 'a', 'b', 'c', 'd', or 'e'.
```
>>> PHYS121.get_hash(a4_7.lower()) == '4a8a08f09d37b73795649038408b5f33'
True
```
failure_message: Try again.


Tests for Q4.8 (points: None)
```
>>> a4_8 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_8, str)
True
```
failure_message: a8 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_8, str):
...     test = a4_8.lower() in ['a', 'b', 't', 'f']
>>> test
True
```
failure_message: a8 must be one of 'a', 'b', 't', or 'f'.
```
>>> PHYS121.get_hash(a4_8.lower()) in ['0cc175b9c0f1b6a831c399e269772661', 'e358efa489f58062f10dd7316b65649e']
True
```
failure_message: Try again.


Tests for Q4.9 (points: None)
```
>>> a4_9 != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(a4_9, str)
True
```
failure_message: a9 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(a4_9, str):
...     test = a4_9.lower() in ['a', 'b', 't', 'f']
>>> test
True
```
failure_message: a9 must be one of 'a', 'b', 't', or 'f'.
```
>>> PHYS121.get_hash(a4_9.lower()) in ['0cc175b9c0f1b6a831c399e269772661', 'e358efa489f58062f10dd7316b65649e']
True
```
failure_message: Try again.


Tests for Q5.1 (points: None)
```
>>> x != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(x, int)
True
```
failure_message: x must be an integer.
```
>>> test = True
>>> if isinstance(x, int):
...     test = len(str(x)) == 6
>>> test
True
```
failure_message: x must contain exactly 6 digits.
```
>>> test = True
>>> if isinstance(x, int):
...     test = str(x)[0:3] == str(x)[3:6]
>>> test
True
```
failure_message: The first 3 digits of x must be identical to the last 3 digits.


Tests for Q5.2 (points: None)
```
>>> y != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(y, int) or isinstance(y, float)
True
```
failure_message: y must be an integer or a decimal number.
```
>>> test = True
>>> if isinstance(y, int) or isinstance(y, float):
...     test = int(str(x)[0:3]) == int(y)
>>> test
True
```
failure_message: Try again.


Tests for Q6.1 (points: None)
```
>>> num != ...
True
```
failure_message: Replace the ... with your answer.
```
>>> isinstance(num, str)
True
```
failure_message: num must be a string (i.e. surrounded by single or double quotes).


Tests for name_and_student_number_1 (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)


Tests for name_and_student_number_2 (points: None)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...
