# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">**PHYS 121 Lab #8**<span style="color:red"> $\to$ (19 possible marks)</span> </span></div> 

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# Faraday's Law \& Eddy Currents

***
## Learning Objectives:
* <b><span style="color:rgba(0, 153, 76, 1);">Quantitatively study and model the motion of a strong magnet sliding down a conducting track.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Observe static magnetic fields exerting a force on moving charges.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Observe changing magnetic fields inducing currents.</span></b>

***
## Autograding:
The PHYS 121 Pre-lab assignments and Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS121.Installer()' reports that some functions have been installed, the user should restart the "kernel" and then exectute the PHYS121.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS121.Installer()' command requires the file 'PHYS121.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Import PHYS121.py and then run the installer function.
import PHYS121
PHYS121.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 121 - Lab 8.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
from IPython.display import YouTubeVideo
import math
pd.set_option('display.max_rows', None) # Display all DataFrame rows
```

```python
grader.check("q0.1")
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Name & Student Number</span></div>  

To ensure that both you and your partner receive credit for you lab work, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number as strings (i.e. place your names and student numbers between single or double quotes.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
PHYS121.graderCheck([last_name_1, first_name_1, student_number_1, last_name_2, first_name_2, student_number_2], ['last_name_1', 'first_name_1', 'student_number_1', 'last_name_2', 'first_name_2', 'student_number_2'], grader.check('name_and_student_number'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Introduction (15 minutes)</span></div>

Electromagnetism can be unintuitive at times - this lab will help demonstrate a few principles from electromagnetism so that even if you think they are *strange*, you can at least clearly observe their effects.  Many of the concepts relevant to this lab have not been covered in the lecture yet.  That's okay.  In research, it's common for experimental work to predate a full theoretical understanding.  For example, superconductivity was first observed in mercury in 1911, but a theoretical understanding of the phenomenon wasn't fully developed until 1957 in the so-called BSC (Bardeen, Cooper, and Schrieffer) Theory of Superconductivity.  A new class of *high-temperature* superconductors was discovered in 1987.  A microscope understanding of the mechanism driving superconductivity in these materials still eludes physicists to this day.

> **Today's lab will be used to motivate the topics that will be discussed in the PHYS 121 lectures over the remaining few weeks in the term.** 😉


Recall that electrical currents are a *source* of magnetic fields.  We will see that **Ampère's law** gives us a way to calculate the magnetic field produced by an arrangement of currents.  In Lab #7, we used the fact that a current in a solenoid produces a uniform magnetic field within the bore of the solenoid.  **Faraday's law** describes a complementary process in which magnetic fields produce electrical currents. Mathematically, this effect is expressed as:

\begin{equation}
\mathcal{E} = -\frac{d\Phi_B}{dt}
\tag{1}
\end{equation}

where $\Phi_B$ is the **magnetic flux** through a closed conducting loop and $\mathcal{E}$ is the **electromotive force** or **emf** produced in the loop.  The term **emf** is just used to represent a voltage *induced* by a changing magnetic flux.  The units of **emf** are volts: $[\mathcal{E}]=\rm\ V$.  Notice from Eq. (1) that, if the magnetic flux $\Phi_B$ is constant, then $\mathcal{E}=0$.  

One important observation is that we now have a second way of producing voltage differences.  Previously, we always had to rely on batteries to provide a voltage difference to drive current in a circuit.  It now appears that a changing magnetic field (or more precisely, magnetic flux) can be used to establish a voltage differnce to drive current in a circuit.  This phenomenon is the basis for electric generators which convert mechanical energy into electrical energy.

A second important observation is that:

 - only *moving* charges create magnetic fields
 - only a *changing* magnetic flux can induce an emf
 
 A charge at rest produces no magetic field and a constant magnetic flux produces no emf.

The magnetic flux through a surface is determined by the component of the magnetic field that is perpendicular to the surface. For flat surfaces, $\Phi_B = BA\cos\theta$, where $A$ is the area of the surface and  $\theta$ is the angle between $\vec{B}$ and a vector normal to the surface. **Figure 1**, shows three surfaces with identical areas $A$ in the same uniform magnetic field $\vec{B}$. 
* Surface A is perpendicular to the field, such that $\theta = 0$, $\cos\theta = 1$, and $\Phi_B = BA$.
* The normal to surface B (the blue vector) is at an angle $\theta$ with respect to $\vec{B}$, such that $\Phi_B = BA\cos\theta = BA_{\perp}$, where $A_{\perp}$ is the area of the projected surface surrounded by dashed lines and $\perp$ to $\vec{B}$ in **Fig. 1**.
* Surface C is parallel to $\vec{B}$ such that its normal vector is perpendicular to the magnetic the field ($\theta = 90^\circ$). Since $\cos 90 = 0$,  $\Phi_B = 0$ in this case.

<br>
<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-8/images/flux.png" style="width: 60%" /><br>  
<b>Fig. 1:  Although the surfaces A, B and C experience the same uniform magnetic field $\vec{B}$ and have the same surface area $A$, the magnetic flux $\Phi_B$ through each depends on their orientation with respect to the magnetic field. </b></center>
</p>

If these surfaces were replaced by conducting loops, we would *still not* expect a current to be generated because the magnetic flux through each loop/surface is constant: the induced emf in Eq. (1) depends on the *time-rate of change* of the $\Phi_B$, and not the flux itself. If these loops/surfaces were fixed in space and $\vec{B}$ was constant, there would be no emf and hence no induced current.

Since, for a flat surface, $\Phi_B = BA\cos\theta$, there are three different ways to make a *changing* magnetic flux.  We could manipulate any of $B$, $A$, or $\theta$:

- We could change the strength of the magnetic field.  For example, as shown in **Fig. 2**, moving a permanent magnet towards or away from a loop of wire will change the magetic flux through the loop which induces and emf that establishes a current in the loop.  Effectively, this is the effect that you'll investigate in today's lab.

<p>
<center>
<img src="https://github.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/blob/main/Lab-8/images/Faraday.jpg?raw=true" style="width: 60%" /><br>  
<b>Fig. 2:  Moving a permanent magnetic towards or away from a coil of wire induces a current in the wire.  Figure taken from <a href = "https://phys.libretexts.org/Courses/Muhlenberg_College/Physics_122%3A_General_Physics_II_%28Collett%29/09%3A_Electromagnetic_Induction/9.02%3A_Faradays_Law">LibreTexts</a>. </b></center>
</p>

- Alternatively, we could change the area of a loop suspended in a uniform magnetic field.  For example, **Fig. 3** shows a square loop of wire in a uniform magnetic field.  If one was to close the loop by pulling opposite corners away from one another, there would be an induced current in the loop during the time it takes area to go from $L^2$ to zero.

<p>
<center>
<img src="https://github.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/blob/main/Lab-8/images/Faraday1.png?raw=true" style="width: 40%" /><br>  
<b>Fig. 3:  Changing the area of a loop in a uniform magnetic field will result in an induced emf and current.</b></center>
</p>

- Finally, as shown in **Fig. 4**, we could also change the angle $\theta$ that the normal to the loop makes with $\vec{B}$.  At the instant the loop is perpendicular $\vec{B}$, the magnetic flux is a maximum and *decreasing*.  When the loop is parallel to $\vec{B}$, the flux is a minimum and *increasing*.  **Figure 4** demonstrates the basic idea behind electric generators which supply nearly all of the power in electric power grids.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/59bb1e708915cfd509d6d150ff0c9f7c1e355ad0/Lab-8/images/Faraday2.svg" style="width: 33%" /><br>  
<b>Fig. 4:  Rotating a loop of fixed area in a uniform magnetic field will result in an induced emf and current.  Figure adapted from an image taken from <a href = "https://commons.wikimedia.org/wiki/File:Electric-generator-3D-with-load.svg">Wikimedia Commons</a>.</b></center>
</p>



# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Magnetic Braking - Data Collection (50 minutes)</span></div>

In the first part of this lab, you will explore one effect of Faraday's Law of induction.  **Figure 5** shows a magnet sliding along a copper plate with velocity $\vec{v}$. Since copper is a conductor, any circle embedded in the copper sheet can be considered a conducting loop.  The sliding magnet, and its magnetic field, are moving away from the loop on the left causing the magnetic flux through the left loop to *decrease*.  On the other hand, the magnet is approaching the loop on the right causing the magnetic flux to *increase* in this loop.  

<br>
<p>
<center>
<img src="https://github.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/blob/main/Lab-8/images/eddy.png?raw=true" style="width: 55%" /><br>
<b>Fig. 5: Schematic diagram of the eddy currents and resulting magnetic fields produced by a magnet sliding along a copper sheet.  </b>
</center>
</p>

Notice that the induced currents in these two loops have been drawn in opposite directions.  Remember also that currents create magnetic fields.  Therefore, we have the following situation:

- A changing magnetic flux induces an emf $\mathcal{E}$
- The induced emfs drive currents $I_\mathrm{ind}$ in the copper strip
- The induced currents create their own magnetic fields $\vec{B}_\mathrm{ind}$

**Lenz's law** states that the induced magnetic fields act to **oppose** the original change in magnetic flux.  In **Fig. 5**, the magnetic flux in the left loop is decreasing.  The original magnetic field due to the sliding magnet was pointing up.  $I_\mathrm{ind}$ in this loop acts to oppose this change in flux and, hence, produces a $\vec{B}_\mathrm{ind}$ that also points *up*.  The situation is the opposite for the right loop:  The magnetic flux is increasing and the induced current creates a $\vec{B}_\mathrm{ind}$ that points *down* to oppose this change.  Little bar magnets have been added to the figure to emphasize the directions of the magnetic fields due to the sliding magnet and the left and right current loops.  The south pole at the bottom of the sliding magnet will be attracted to the north pole of the left current loop and repelled by the south pole of the right current loop.  As a result, even in the absence of friction, the sliding magnet will slow down and eventually come to a stop, a demostration of *magnetic braking*.

We will attempt to model the magnetic braking force as a drag force.  Specifically, we assume that the braking/drag force $F_\mathrm{d}$ will be proportional to the speed of the magnet.  This assumption makes sense because, if the magnet is moving fast, the magntic flux will be changing quickly (large $d\Phi_B/dt$) and the induced emf will be large.  Thefore:

\begin{equation}
F_{d} = -bv
\tag{2}
\end{equation}

where $v$ is the velocity and $b$ is a drag coefficient.  Recall that a similar model was used for the drag force in PHYS 111 when you investigated the drag acting on a coffee filter dropped from some height above the ground.  Of course, there will also be kinetic friction between the magnet and copper track.  **Figure 6** shows a complete free-body diagram of the forces acting on the magnet as it slides down an inclined ramp.  $\mu_\mathrm{k}$ is the coefficient of kinetic friction, and $\vec{N}$ is the normal force.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-8/images/FBD.png" style="width: 40%" /><br>
<b>Fig. 6: Free-body diagram for the magnet sliding down an inclined copper ramp.  </b></center>
</p>

Applying Newton's 2$^\mathrm{nd}$ law to this system leads to $N = mg\cos\alpha$ and:

$$
ma = mg\sin\alpha -\mu_\mathrm{k}mg\cos\alpha -bv
\tag{3}
$$

where $a$ is the magnitude of the acceleration down the ramp.  If the magnet is released from rest, its speed will initially increase until it reaches a terminal velocity $v_\mathrm{T}$ at which point $a=0$.  Setting $a=0$ in Eq. (3) leads to:

$$
v_\mathrm{T} = \frac{mg}{b}\left(\sin\alpha - \mu_\mathrm{k}\cos\alpha\right)
\tag{4}
$$

Substituting this expression for $v_\mathrm{T}$ back into Eq. (3) lead to:

$$
ma = b\left(v_\mathrm{T} - v\right)
\tag{5}
$$

By recognizing that $a=\dfrac{dv}{dt}$ and $v=\dfrac{dx}{dt}$, it is possible to solve Eq. (5) for the speed $v$ and position $x$ of the sliding magnet as a function of time.  The results are:

\begin{align}
v(t) &= v_\mathrm{T}\left(1-e^{-t/\tau}\right)\tag{6}\\
x(t) & = v_\mathrm{T}\tau\left[\dfrac{t}{\tau} - 1 + e^{-t/\tau}\right]\tag{7}
\end{align}

where $\tau = m/b$ is the time constant that determines how quickly the magnet approaches the terminal velocity.


<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.1</span>**  **<span style="color:red">(1 mark)</span>**  

What would you expect a graph of $x(t)$ just after the magnet is released and it's *approaching* terminal velocity? If you're having trouble working this out from Eq. (7), try using your intuition, and then refer back to the equation to make your statement more precise.

***
**<span style="color:blue">Answer 2.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.2</span>**  **<span style="color:red">(1 mark)</span>**  

Describe the expected shape of the graph of $x(t)$ *after* the magnet has reached terminal velocity.

***
**<span style="color:blue">Answer 2.2:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 
***

<!-- END QUESTION -->

On your table, you'll find magnet and a ramp with three copper tracks. There's also enough space between two of the copper tracks to slide the magnet down a plastic track with no copper beneath it.  The three copper stips have different thicknesses.  Referring to **Fig. 7**, from left to right, the thickness of the copper tracks are:
* Thin  $\Longrightarrow 0.006\rm\ inches = 0.152\rm\ mm$ (copper foil)
* Medium $\Longrightarrow 0.021\rm\ inches = 0.533\rm\ mm$
* Thick $\Longrightarrow 0.0625\rm\ inches = 1.59\rm\ mm$

In **Fig. 7**, the plastic track is between the thin and medium copper tracks.  For some of the ramps in the lab, the plastic track is between the medium and thick copper tracks. 

<p>
<center>
<img src="https://github.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/blob/main/Lab-8/images/ramps.png?raw=true" style="width: 65%" /><br>
<b>Fig. 7: Materials for Part 2 of this lab. Left: Top and bottom of the permanent magnet. The top has a dot to facilitate video tracking, and the bottom is covered with Teflon to reduce kinetic friction. Right: Ramp with three copper tracks and one plastic track.</b></center>
</p>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.3</span>**  **<span style="color:red">(1 mark)</span>**  

We have introduced an unknown parameter (the drag coefficient) $b$ into our expression for the drag force. The value of $b$ sets the strength of the magnetic braking effect - if $b$ is large, the effect will be strong. Rank the four tracks (three copper and one plastic) in terms of the expected value of $b$.  Order your ranking from the largest value of $b$ to the smallest. Briefly explain your choice.

***
**<span style="color:blue">Answer 2.3:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 
***

<!-- END QUESTION -->

Using the provided phone, you will record the magnet sliding down each of these tracks. **In order to accurately track the magnet's motion, you will need to adjust the camera to its slow-motion setting.**  When in the slow-motion mode, you won't be able to adjust the zoom setting of the camera.  Instead, you'll need to adjust the height that the phone is held above the track to a suitable distance.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-8/images/slomo.jpg" style="width: 25%" /><br>
<b>Fig. 8: With the camera app opened, scroll to the "More" tab to switch to the slow motion setting.  </b></center>
</p>

Before filming, ensure that the phone is oriented parallel to the ramp and that you've maximized the region of the ramp in view of the camera. You wil record four separate videos and then upload them to the desktop computer at your station. The following steps will guide you through using Tracker to collect data from your videos. 

1. Move your videos from the DCIM folder to the home drive on the computer (H:)

2. Open Tracker. 

3. For each of your recordings, complete steps 3 - 12 below.  It is best if you keep the videos short.  It will save time when importing the recordings into Tracker.

4. Click File -> Import -> Video... and navigate to the appropriate folder to import your video

5. The video will may be oriented in the wrong way. To fix this, Right-click on video -> Filters -> New -> Rotate 

6. To set the frame-rate, right click Video -> Clip Settings... and change the framerate from $30\rm\ FPS$ to $240\rm\ FPS$.

7. Find the first frame of the video after the magnet was released. You can click and drag the scrollbar or use Ctrl+mouse wheel to move through the frames. Take note of the frame number, and then move the left-hand black arrow to this frame. 

8. Do the same thing for the indicator of the ending time. Tracker will only consider the frames between the start and stop indicators. Click the black return arrow on the left side of the scroll bar to return to the first marked frame.  See **Fig. 9**.  **Note: Your results will be best if you track the position of the magnet for only the first $0.3$ to $0.4\rm\ s$ of its motion after it is released.**  This corresponds to about 100 frames of the video clip.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-8/images/scrollbar.png" style="width: 60%" /><br>
<b>Fig. 9: Scrollbar before and after steps 7 and 8. The black arrows on the bottom bound the segment that you want to measure.  </b></center>
</p>

9. Place your axes so that the origin is at the centre of your magnet in the first frame. Rotate the axes so that the magnet will be sliding in the positive x-direction.

10. To set the length scale, press the calibration icon, then New -> Calibration Stick. Click and drag the calibration stick to align it with the orange ruler on your ramp and set the distance to $0.150\rm\ m$.  See **Fig. 10**. (For added precision, use the mouse wheel to zoom in and out.)

<p>
<center>
<img src="https://github.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/blob/main/Lab-8/images/axes_scale.jpg?raw=true" style="width: 35%" /><br>
<b>Fig. 10: Guide to the icons used in steps 9, and 10.  </b></center>
</p>

10. Now set up your tracker. Click Track -> New -> Point Mass and then open the auto-tracker. On the first frame, Shift-Ctrl-Click the centre of your magnet. The inner circle is the "template" - this is what the auto-tracker will look for as the frames advance. It should be large enough to show a clear distinction between the black dot and the pink paper. The outer square is the area in which the auto-tracker will look for a match to the initial template. If your magnet is going very fast by the end of the video, or if the lighting changes abruptly, autotracker may lose sight of the dot. If the suggested mark is accurate, you can press "Accept" to continue. If it isn't, you can manually mark the problem frames using Shift+Click or define a new template entirely using Shift+Ctrl+Click. Another option is to increase the "Evolve" value: Tracker updates the template after each new frame by superposing the old template with the new matches. By increasing "Evolve", you permit more drastic updates. Be careful not to do this unless absolutely necessary, as it may cause inaccuracies. See **Fig. 11**.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-8/images/autotracker.png" style="width: 60%" /><br>
<b>Fig. 11: To set up the autotracker, you'll need to define both the template and the search region.  </b></center>
</p>


11. Click File -> Export -> Data... and change "Number Format" to "Full Precision", "Delimiter" to "Comma", and "Columns" to "t,x" (**make sure you had your axes set appropriately at the beginning of the video! It's not too late to rotate them so that the trajectory is along the positive x-axis**). Once you've exported the .txt file, move it to the same folder as this notebook (drag and drop the file into the frame along the left-hand side of the notebook).  In order for your work to be graded properly, you must name your exported files **thin.txt**, **medium.txt**, **thick.txt**, and **plastic.txt**. 

12. Click File -> New Tab and repeat.

In case it's helpful, here are links to the Tracker YouTube videos that were included in Pre-Lab #8

 - <a href = "https://youtu.be/MrAZ1PbB_ek">Importing and rotating videos</a>
 - <a href = "https://youtu.be/6cvNUosIac0">Calibrating the length scale</a>
 - <a href = "https://youtu.be/XlbKkyvWlr4">Setting the origin and orientation of the coordinate axes</a>
 - <a href = "https://youtu.be/O5_AMa4OL9s">Setting the start and stop limits</a>
 - <a href = "https://youtu.be/6FhGTUuTIMY">Setting up and using the auto-track feature</a>
 - <a href = "https://youtu.be/kxEvWegEKOM">Exporting the data</a>



# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Magnetic Braking - Data Analysis (30 minutes)</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.1</span>**  **<span style="color:red">(4 marks)</span>**  

The first cell below contains a template for loading your **thin.txt** file into a DataFrame. Follow this template to upload the rest of your data.  You just need to change the file name ```"thin.txt"``` appropriately.

*** Please do not change anything to the left of the equals sign. ***

```python
# Import data for the thin copper track.
thin = pd.read_csv("thin.txt", skiprows=2, header = None, names=["t","x"])

# Display the dataFrame.
thin
```

```python
# Import data for medium copper track
medium = ...

# Display the dataFrame.
medium
```

```python
# Import data for the thick copper track
thick = ...

# Display the dataFrame.
thick
```

```python
# Import data for plastic track
plastic = ...

# Display the dataFrame.
plastic
```

<!-- END QUESTION -->

Consider Eq. (7) once more, which is repeated below. It shows exponential behaviour in the approach to terminal velocity:

\begin{equation}
x(t) = v_\mathrm{T}\tau\left( \frac{t}{\tau} - 1 + \exp ( -t/\tau ) \right)
\tag{8}
\end{equation}

It's common practice to consider *limiting cases*. What happens when our parameter $\tau = m/b$ is very large or very small? Recall that:

\begin{equation}
v_\mathrm{T}  = \tau g\left( \sin\alpha - \mu_\mathrm{k} \cos\alpha \right) 
\end{equation}

In the limit $\tau \to \infty$ ($m >> b$ or $b \to 0$) it is possible to show that $x(t)\propto t^2$, which is what we expect for an object experiencing a constant acceleration due to gravity  and kinetic friction:

\begin{equation}
\lim\limits_{\tau\to\infty}x(t) \approx \frac{v_\mathrm{T}}{2\tau}t^2 = \frac{1}{2} gt^2 \left( \sin \alpha - \mu_\mathrm{k} \cos \alpha \right)
\tag{9}
\end{equation}

At the opposite extreme, as $\tau \to 0$ ($b >> m$ or $b \to \infty$) we see that $x(t)$ is linear, suggesting that terminal velocity is reached almost instantly when the magnetic braking is strong:

\begin{equation}
\lim\limits_{\tau\to 0}x(t) = v_\mathrm{T}t
\tag{10}
\end{equation}

## **<span style="color:blue">Summary</span>**

Recall that $b$ is the drag coefficient that slows the magnet as it slides down the track.
 - In the $b\to 0$ limit, there is no magnetic braking and $x(t)\propto t^2$
 - In the large-$b$ limit, the magnetic braking is strong, the magnet quickly reaches terminal velocity, and $x(t)\propto t$
 - In the intermediate regime, $x(t)$ is given by Eq. (8)

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.2</span>**  **<span style="color:red">(2 marks)</span>**   

Answer the following questions based on the discussion above and your observations during the experiment. As always, explain your answers. 


- (a) For which tracks (if any) do you expect Eq. (9) to be appropriate? 
 
- (b) For which tracks (if any) do you expect Eq. (10) to be appropriate?

- (c) For which tracks (if any) do you expect *neither* approximation to be appropriate?

***
**<span style="color:blue">Answer 3.2:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 
***

<!-- END QUESTION -->

Below, we provide three templates for you to use when analyzing your data: a linear fit, a power-law fit, and a fit designed specifically for Eq. (8).

```python
# Linear fit template. 
m, b, dm, db, fig = PHYS121.LinearFit(xData = ...,
                                      yData = ..., 
                                      xlabel = 'time', 
                                      ylabel = 'position', 
                                      xUnits = 's', 
                                      yUnits = 'm')

# Power law fit template.
A, N, C, dA, dN, dC, fig = PHYS121.PowerLaw(xData = ..., 
                                            yData = ..., 
                                            xlabel = 'time', 
                                            ylabel = 'position', 
                                            xUnits = 's', 
                                            yUnits = 'm')

# Eq. (7) fit template.
v_T, tau, dv_T, dtau, fig = PHYS121.Braking(xData = ..., 
                                            yData = ..., 
                                            xlabel = 'time', 
                                            ylabel = 'position', 
                                            xUnits = 's', 
                                            yUnits = 'm')
```

Remember, the access a specific column from a dataFrame called df, you use the syntax:
```python
df["column_name"]
```

For the following four questoins, use the templates above to perform the appropriate fit for each of your four datasets. You're welcome to try multiple fits on a dataset and take the one that provides the *simplest* result without sacrificing accuracy.  If you find that your results contradict your answers from **Question 3.2**, comment on this fact in the question that follows the fits.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.3</span>**  **<span style="color:red">(1 mark)</span>**  

Complete the fit for the thin copper track.

```python
# Thin copper track
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.4</span>**  **<span style="color:red">(1 mark)</span>**  

Complete the fit for the medium copper track.

```python
# Medium copper track.
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.5</span>**  **<span style="color:red">(1 mark)</span>**  

Complete the fit for the thick copper track.

```python
# Thick copper track
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.6</span>**  **<span style="color:red">(1 mark)</span>**  

Complete the fit for the plastic track.

```python
# Plastic track.
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.7</span>**  **<span style="color:red">(2 marks)</span>**  

(a) Comment on the quality of your fits.  Did your fits match your expectations from **Question 3.2**?

(b) Based on your results, what can you say about the relative value of $b$ for each track? How does this compare to your answer from **Question 2.3**?

***
**<span style="color:blue">Answer 3.7:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 
***

<!-- END QUESTION -->

At the back of the room, we have set up a track made of actual coils of wire.  These coils are connected to light-emitting diodes (LEDs).  When current is induced in the coil in one direction a green LED lights up.  When current is induced in the opposite direction, a red LED lights up.  We won't ask you any questions about this demo, but feel free to play with the set up.  The magnet moves pretty quickly which makes it difficult to see precisely when the LEDs turn off and on.  Executing the cell below will open a YouTube video that has slow-motion recordings of the magnet's motion and LEDs turning off and on.

```python
YouTubeVideo("7qMgMOK3uec", width=500)
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - Cathode Ray Tubes (20 minutes)</span></div>

Prior to the invention of liquid crystal displays (LCDs), **cathode ray tubes** (or **CRT**s) were the primary component in television sets. These simple devices were once ubiquitous - in fact, the first invention that might be recognized as a video game was the *cathode ray amusement device*, [patented in 1948 by Thomas Goldsmith](https://www.popularmechanics.com/culture/gaming/a20129/the-very-first-video-game/). CRTs also have a profound legacy within physics -- experiments with cathode ray tubes led to the 1897 discovery of the electron by JJ Thomson. This discovery and a few other puzzling experimental and theoretical developments from the late 19th century would eventually culminate in the quantum revolution of the 1920s. 

**Figure 12** shows a simplified diagram of a cathode ray tube. A tungsten heating element (much like the filament of an incandescent lightbulb) is used to heat up a nearby cathode, typically made from barium carbonate. A voltage source creates a slight excess of electrons in the cathode. The kinetic energy of these electrons increases with temperature until the most energetic electrons are released from the cathode in a process called **thermionic emission**.

The emitted cloud of electrons is concentrated into a beam and accelerated by an electric field. Magnetic fields established by deflecting coils are used to steer the beam as it is fired towards a phosphorescent screen which begins to glow a dull green. 

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-8/images/crt.png" style="width: 65%" /><br>
<b>Fig. 12: Simplified diagram of a cathode ray tube. Note that phosphorence is the same phenomenon exploited by glow-in-the-dark materials.  </b></center>
</p>

There will be stations in the lab room with CRTs and bar magnets. You are free to experiment with these, and you should attempt to predict what will happen before you try different things. Recall that the force on a moving charge is given by:

\begin{equation}
\vec{F}_B = q \vec{v} \times \vec{B} 
\tag{8}
\end{equation}

for a charge $q$ moving with velocity $\vec{v}$ through a magnetic field $\vec{B}$. Recall also that the **right-hand rule** tells you how to determine the direction of the cross product of two vectors:

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-8/images/rhr.png" style="width: 35%" /><br>
<b>Fig. 10: The right-hand rule for determining the direction of $\vec{a}\times\vec{b}$, the cross product of vectors $\vec{a}$ and $\vec{b}$. From <a href="https://commons.wikimedia.org/wiki/File:Right_hand_rule_cross_product.svg">Wikimedia Commons</a>, licenced under <a href="https://creativecommons.org/licenses/by-sa/3.0/legalcode" rel="license">CC BY-SA 3.0</a>.  </b>
</center>
</p>


The following questions are purely qualitative and require no calculations, but you still must explain as specifically as possible *why* one would expect the observed behaviour. 

#### <span style="color:purple">The 'PHYS121.graderCheck(...)' statements for **Questions 4.1** and **4.2**,  only check that answers have been entered in the correct format.</span>  
#### <span style="color:purple">They do **NOT** check that you've entered the correct answer.</span>
#### <span style="color:purple">You can check your answers by doing the experiments at the back of the room!</span> 

***
**<span style="color:blue">Question 4.1</span>**  **<span style="color:red">(1 mark)</span>**  

Bring the bar magnet towards the CRT as shown:

<br>
<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-8/images/crtmag1.png" style="width: 19%" />
</center>
</p>

What (if anything) will happen to the beam of electrons, as shown by the dot on the phosphorescent screen?  The *negative* beam of electrons is moving out of the screen.  Also remember that magentic field lines exit north poles and enter south poles.

- (a) The beam will be deflected upwards
- (b) The beam will be deflected downwards
- (c) The beam will be deflected to the left
- (d) The beam will be deflected to the right
- (e) The beam will **not** be deflected

***
**<span style="color:blue">Answer 4.1:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, or e) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_1 = ...
```

```python
PHYS121.graderCheck([a4_1], ['a4_1'], grader.check('q4.1'))
```

***
**<span style="color:blue">Question 4.2</span>**  **<span style="color:red">(1 mark)</span>**  

<br>
<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-8/images/crtmag2.png" style="width: 25%" />
</center>
</p>

If you bring the bar magnet in from the side (as in the figure above) what happens? (Hint: try picturing the magnetic field lines around the bar magnet) 

- (a) The beam will be deflected upwards
- (b) The beam will be deflected downwards
- (c) The beam will be deflected to the left
- (d) The beam will be deflected to the right
- (e) The beam will **not** be deflected

***
**<span style="color:blue">Answer 4.2:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, or e) between single or double quotes.  
*** Please do not change anything to the left of the equals sign. ***

```python
a4_2 = ...
```

```python
PHYS121.graderCheck([a4_2], ['a4_2'], grader.check('q4.2'))
```

***
Finally, here's another short YouTube video.  It is a recording from the same CRT used in the lab.  Watch what happens when a rotating magnetic field is placed near the electron beam.  Again, we won't ask any more questions, just enjoy the intricate interplay between electrcity and magnetism!

```python
YouTubeVideo("UgMZC8fT0zM", width=500)
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 5- Statement of contributions (10 minutes)</span></div>

***
Below, you will briefly summarize the contributions that you and your partner made when completely this lab.  Please ensure that both partners have a chance to engage in all types of lab activities (developing the experimental methods/data collection strategies, hands-on interaction with the apparatus/data collection, and data processing/analysis in the Jupyter notebook).  Note that statements of contributions are common in science.  The figure below shows one example taken from a 2021 paper published in the journal *Nature Physics*.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/contributions.jpg" alt="image info" style="width: 60%; height: auto;" /><br>
<b>Fig. 4: A sample statement of author contributions from a recent journal article.</b></center>
</p>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 5.1(a):</span>**  

Answer this question only if you worked with a partner to complete this lab.

In two or three sentences describe the specific contributions of **Partner #1** to this lab.  For example, worked with the experimental apparatus to collect data for Part 1 of the lab, entered and processed the data in the Jupyter notebook for Part 2 of the lab, ...  

***
**<span style="color:blue">Answer 5.1(a):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  Please start be entering the name and student number of **Partner #1**.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 5.1(b):</span>**  

Answer this question only if you worked with a partner to complete this lab.

In two or three sentences describe the specific contributions of **Partner #2** to this lab.  

***
**<span style="color:blue">Answer 5.1(b):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  Please start be entering the name and student number of **Partner #2**.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 6 - Student Surveys (20 minutes)</span></div>

The last thing that we'd like to ask you to do in the PHYS 121 Lab is complete two online surveys.  You responses will help us evaluate what's working well in the PHYS 121 labs and areas that are in need of improvement.

### Survey #3 - Experience in the Lab (End-of-Term)

The first survey is used to collect feedback about the lab activities and the Jupyter notebook-based lab manuals.  The lab activities and method of delivery are both new, so we really do value your feedback.

Here is the link to [Survey #3 - Experience in the Lab (End-of-Term)](https://ubc.ca1.qualtrics.com/jfe/form/SV_2o8U4lBTz74KDmS)


### Survey #4 - TA Quality (End-of-Term)

The second survey is used to collect feedback about the quality of your TA.  The anonymous feedback will be shared with your TA so that in future courses they can keep doing the things that you like while simultaneously trying to improve any weaknesses. 

Here is the link to [Survey #4 - TA Quality (End-of-Term)](https://ubc.ca1.qualtrics.com/jfe/form/SV_9M53F8dCITGCVVk)

Both of these surveys are also available in the **PHYS 121 Labs 2022W2** [Canvas](https://canvas.ubc.ca/) shell.  In Canvas, click on "Modules" and then scroll down to the bottom of the page to find the "Feedback" module. If your partner is using the lab computer to complete the surveys, feel free to log into Canvas to complete the two surveys.

***
**<span style="color:blue">Question 6.1</span>** **<span style="color:red">(1 mark)</span>**

Partner #1 and Partner #2 (if applicable), completed both of the online surveys.

(a) True  
(b) False

***
**<span style="color:blue">Answer 6.1:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, T, or F) between single or double quotes.  

*** Please do not change anything to the left of the equals sign. ***

```python
a6_1 = ...
```

```python
PHYS121.graderCheck([a6_1], ['a6_1'], grader.check('q6.1'))
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 7 - Submission (5 minutes)</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 7.1:</span>**  

We welcome your feedback on the PHYS 121 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 121 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 7.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file using the link that appears **below** the ```grader.check_all(...)``` cell.
- $\Leftarrow$ Please do **NOT** download the zip file from the frame along the left-hand side of the screen. 
- Upload the .zip file to the PHYS 121 Canvas gradebook.
- **Do NOT change the name of the .zip file.**
- **Do NOT modify the contents of the .zip file.**

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
display(grader.check_all())
PHYS121.saveMessage()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
PHYS121.waitMessage()
grader.export(files=["thin.txt", "medium.txt", "thick.txt", "plastic.txt", "PHYS121_DataLogger.txt"])
PHYS121.proceedMessage()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 8 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.  
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.  
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: March 23, 2024

Tests for name_and_student_number (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...


Tests for q0.1 (points: None)
```
>>> import sys
>>> 'PHYS121' in sys.modules
True
```
failure_message: Make sure you execute the cells containing "import PHYS121" and "import otter" as well as the cell beneath the heading "Import Modules".


Tests for q4.1 (points: None)
```
>>> isinstance(a4_1, str)
True
```
failure_message: Your answer must be a string, i.e. surrounded by single or double quotes.
```
>>> a4_1.lower() in ['a', 'b', 'c', 'd', 'e']
True
```
failure_message: Your answer must be a string of one of 'a', 'b', 'c', 'd', or 'e'.


Tests for q4.2 (points: None)
```
>>> isinstance(a4_2, str)
True
```
failure_message: Your answer must be a string, i.e. surrounded by single or double quotes.
```
>>> a4_2.lower() in ['a', 'b', 'c', 'd', 'e']
True
```
failure_message: Your answer must be a string of one of 'a', 'b', 'c', 'd', or 'e'.


Tests for q6.1 (points: None)
```
>>> a6_1.lower() in ['a', 'b', 't', 'f']
True
```
failure_message: Your answer must be a string of one of 'a', 'b', 'T', or 'F'.
