# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">**PHYS 121 Lab #5**<span style="color:red"> $\to$ (24 possible marks)</span> </span></div> 

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# Electric & Hydraulic Circuits &ndash; Week 2

***
## Learning Objectives:
* <b><span style="color:rgba(0, 153, 76, 1);">Use concepts and mathematical tools developed to solve a problem in one field of physics to solve seemingly unrelated problems from other fields of physics.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Identify the limitations of analogies between distinct physical systems.</span></b>

In this week's lab, you will first use the data from Lab #4 to analyze the dependence of hydraulic resistance $R_\mathrm{H}$ on the inverse of the inner diameter $d^{-1}$ of the Teflon tubes that were used.  Next, you the same hydraulic resistors (small-diameter Teflon tubes) and hydraulic batteries (large vessel filled with water) from Lab #4 to charge a hydraulic capacitor (a large-diameter glass tube).  This experiment is the hydraulic analogue of charging a capacitor using a battery and a resistor.

***
## Autograding:
The PHYS 121 Pre-lab assignments and Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS121.Installer()' reports that some functions have been installed, the user should restart the "kernel" and then exectute the PHYS121.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS121.Installer()' command requires the file 'PHYS121.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Import PHYS121.py and then run the installer function.
import PHYS121
PHYS121.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 121 - Lab 5.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
import pandas as pd
import scipy.stats as sp
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
from IPython.display import YouTubeVideo
import math
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Name & Student Number</span></div>  

To ensure that both you and your partner receive credit for you lab work, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number as strings (i.e. place your names and student numbers between single or double quotes.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
PHYS121.graderCheck([last_name_1, first_name_1, student_number_1, last_name_2, first_name_2, student_number_2], ['last_name_1', 'first_name_1', 'student_number_1', 'last_name_2', 'first_name_2', 'student_number_2'], grader.check('name_and_student_number'))
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - $R_\mathrm{H}$ vs $1/d$ (30 minutes)</span></div>

The first part of this lab will be used to complete the analysis of hydraulic resistance $R_\mathrm{H}$ data that you and your classmates collected in Lab #4. Here is the [link to the Google Sheets document](https://docs.google.com/spreadsheets/d/1BCeJboA1V6BY9LYW8Yyb8QwEOUXmWDtQglrZTukf02c/edit?usp=sharing) that was used at the end of Lab #4.  The first 10 tabs (corresponding to lab sections 'L01', 'L02', ..., 'L10') contain data for Teflon tubes with inner diameters of $d = 0.965$, $1.397$, $1.676$, $1.930$, and $2.413\rm\ mm$.  There is also a 'Physics Labs Team' tab (the last tab) that contains the hydraulic resistance data for narrow Teflon tubes with inner diameters of $d = 0.305$ and $0.559\rm\ mm$. The data from the Google Sheets document has already been added to a csv (**c**omma-**s**eparated-**v**alues) file called "all-RH-Data.csv" that should be among the files listed along the left-hand side of the screen.

### *Electrical Resistance $R$ versus $1/d$*

Recall from **Part 1** from Lab #4 that you measured the electrical resistance of wires having eight different diameters.  Figure 1 below shows a typical plot of the $R$ versus $1/d$ data from Lab #4.  These data were fit to a power law of the form:
$$
R = A\left(\frac{1}{d}\right)^N + C
$$
and returned a value of $N \approx 2$ which suggests that, for electrical resistors, $R\propto d^{-2}$.  It is also worth noting that the constant $C$ returned a value that is essentially zero within experimental uncertainty.


<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-5/images/R-vs-d_inv.png" alt="image info" style="width: 35%; height: auto;" /><br>
<b>Fig. 1: Typical electrical resistance versus $1/d$ data from Lab #4.  The data were fit to a power-law function $R = A(1/d)^N + C$ and returned values of $N \approx 2$ and $C\approx 0$.</b></center>
</p>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.1</span>** **<span style="color:red">(1 mark)</span>** 

The first task is to plot $R_\mathrm{H}$ as a function of $1/d$ and then fit it to the same power-law function.  Begin by executing the cell below to import the data in "all-RH-Data.csv" into a DataFrame called 'Lab4Data'.  The data have been sorted by increasing diameter.

**Note:** Data points from the Google Sheets document that did not follow the trend of the rest of the data have already been excluded from "all-RH-Data.csv".  

*** Please do not change any of the variable names in the cell below ***

```python
# By default, the DataFrame will display only a finite number of rows.
# We can override this default by changing the 'max_rows' option to "None".
# This will force the DataFrame to display all rows.
pd.set_option('display.max_rows', None)

# Load the data from the csv file.
Lab4Data = pd.read_csv("all-RH-Data.csv")

# Display the dataframe.
Lab4Data
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.2</span>** **<span style="color:red">(1 mark)</span>** 

Next, use the 'Scatter()' function from PHYS121.py to plot $R_\mathrm{H}$ as a function of $1/d$.  Include uncertainties in $R_\mathrm{H}$. 

Fill in the missing data needed when calling the 'Scatter()' function.  Remember that the scatter function accepts arrays or lists as inputs - to call a column from a DataFrame as a list, use the syntax:
```python
df["Column Name"] 
```
where 'df' is the name of the DataFrame in question. For additional information about the 'Scatter()' function, see **Appendix A** near the end of this notebook.   
 
 *** Please do not change any of the variable names in the cell below ***

```python
fig1 = PHYS121.Scatter(xData = ..., 
                        yData = ..., 
                        yErrors = ..., 
                        xlabel = "1/d", 
                        ylabel = "R_H", 
                        xUnits = "1/mm", 
                        yUnits = "kg/m^4s")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.3</span>** **<span style="color:red">(1 mark)</span>** 

Next, use the 'PowerLaw()' function from PHYS121.py to plot $R_\mathrm{H}$ as a function of $1/d$ and fit the data to a power-law.  Include uncertainties in $R_\mathrm{H}$. 

The 'PowerLaw()' function is called as follows:
```python
import PHYS121
PHYS121.PowerLaw(xData = x, yData = y, yErrors = erry, xlabel = 'x', ylabel = 'y', xUnits ='s', yUnits = 'm', start = (A0, N0, C0))
```
The values listed in 'start' are used as initial guesses for the expected best-fit parameters.  For additional information about the 'PowerLaw' function, see **Appendix B** near the end of this notebook. 

*** Please do not change any of the variable names in the cell below ***

```python
PHYS121.PowerLaw(xData = ..., 
                yData = ..., 
                yErrors = ..., 
                xlabel = "1/d", 
                ylabel = "R_H", 
                xUnits = "1/mm", 
                yUnits = "kg/m^4s",
                start = (1, 4, 0)
                );
```

<!-- END QUESTION -->

***
**<span style="color:blue">Question 1.4</span>** **<span style="color:red">(1 mark)</span>**  

Based on your experimental determination of $N\pm \Delta N$, what would happen to the hydraulic resistance of a Teflon tube if its inner diameter was doubled?  Pick the closest option.

(a) The resistance would increase by a factor of 16.  
(b) The resistance would increase by a factor of 8.  
(c) The resistance would increase by a factor of 4.  
(d) The resistance would increase by a factor of 2.  
(e) The resistance would remain unchanged.  
(f) The resistance would decrease by a factor of 2.  
(g) The resistance would decrease by a factor of 4.  
(h) The resistance would decrease by a factor of 8.  
(i) The resistance would decrease by a factor of 16.  

***
**<span style="color:blue">Answer 1.4:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, e, f, g, h, or i) between single or double quotes.  

For this question, grader.check() only checks that a valid response has been entered.  It does **not** check if the correct response has been entered. 

*** Please do not change anything to the left of the equals sign. ***

```python
a1_4 = ...
```

```python
PHYS121.graderCheck([a1_4], ['a1_4'], grader.check('q1.4'))
```

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.5</span>** **<span style="color:red">(2 marks)</span>**   

Examine the plot and fit from **Question 1.3** for hydraulic resistance ($R_\mathrm{H}$ vs $1/d$) and compare it to Fig. 1, which is the equivalent plot and fit for electric resistance ($R$ vs $1/d$).  Identify at least two differences between these plots/fits.

***
**<span style="color:blue">Answer 1.5:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - $\log(R_\mathrm{H})$ vs $\log(1/d)$ (30 minutes)</span></div>


From the plot in **Question 1.2**, you may have noticed that its difficult to resolve how $R_\mathrm{H}$ varies for small values of $1/d$.  This happens because the hydraulic resistance at the largest values of $d$ are more than $1000$ times smaller than the corresponding values at the smallest diameters. As we saw in Pre-Lab #4, when plotting data that span such a large range of values (i.e. many orders of magnitude), it is often convenient to view the data using logarithmic scales.

Recall that, in a regular linear plot, the difference in $x$ or $y$ values between neighbouring tick marks along an axis are always equal.  Figure 2 below reproduces plot of $R$ versus $1/d$ for electrical resistance.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-5/images/R-vs-d_inv.png" alt="image info" style="width: 35%; height: auto;" /><br>
<b>Fig. 2: Typical electrical resistance versus $1/d$ data from Lab #4 plotted on a linear scale.</b></center>
</p>

The distance between neighbouring tick marks along the $x$-axis is always $1\rm\ mm^{-1}$ and the distance between neighbouring tick marks along the $y$-axis is always $5\rm\ \Omega$.  

On a plot of $\log y$ versus $\log x$, it is the ratio of neighbouring $x$ and $y$ values that remains fixed.  For example, Fig. 3 below uses the same data from Fig. 2, but plots $\log(R)$ versus $\log(1/d)$.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-5/images/logR-vs-logd_inv.png" alt="image info" style="width: 35%; height: auto;" /><br>
<b>Fig. 3: The same data from Fig. 2 plotted as $\log(R)$ versus $\log(1/d)$.</b></center>
</p>

This time, the difference between tick points along the vertical axis is always $0.2$.  This implies that:

\begin{align}
\log(R_2) - \log(R_1) = 0.2
\label{eq:difference} \tag{1}
\end{align}

where $R_2$ and $R_1$ correspond to the resistance values at any two neighbouring tick marks along the vertical axis.  For this argument, we assume that $R_2 > R_1$.  For example, $\log(R_1)$ and $\log(R_2)$ could correspond to $0.4$ and $0.6$, respectively.  Because $\log A - \log B = \log(A/B)$, Eq. (1) can be re-expressed as:

\begin{align}
\log\left(\frac{R_2}{R_1}\right) = 0.2
\label{eq:log} \tag{2}
\end{align}

Solving Eq. (2) for $R_2/R_1$ yields:

\begin{align}
\frac{R_2}{R_1} = 10^{0.2} = 1.585
\label{eq:ratio} \tag{3}
\end{align}

Thus, in Fig. 3, each tick along the vertical axes corresponds to a change in resistance by a factor of $1.585$.  Advancing by $N$ tick marks corresponds to a change in resistance by a factor of $(1.585)^N$.

***
**<span style="color:blue">Question 2.1</span>** **<span style="color:red">(1 mark)</span>**  

In Fig. 3, the tick makes along the horizontal axis imply that $\log(1/d_2) - \log(1/d_1) = 0.1$.  Determine the ratio of $d_2/d_1$.

(a) $d_2/d_1 = 0.100$  
(b) $d_2/d_1 = 0.631$  
(c) $d_2/d_1 = 0.794$  
(d) $d_2/d_1 = 1.26$   
(e) $d_2/d_1 = 1.59$    
(f) $d_2/d_1 = 10.0$

***
**<span style="color:blue">Answer 2.1:</span>**

Replace the ... in the cell below with your answer.  Your answer should be a single character (a, b, c, d, e, or f) between single or double quotes.  

For this question, grader.check() only checks that a valid response has been entered.  It does **not** check if the correct response has been entered. 


*** Please do not change anything to the left of the equals sign. ***

```python
a2_1 = ...
```

```python
PHYS121.graderCheck([a2_1], ['a2_1'], grader.check('q2.1'))
```

***
Next, go back and compare Figs. 2 and 3 which display identical datasets.  Plotted on a linear scale (Fig. 2), the $R$ versus $1/d$ data exhibit a quadratic dependence.  On the other hand, when plotted as $\log(R)$ versus $\log(1/d)$, the data look linear (another observation from Pre-Lab #4).

The fit in Fig. 2 implies that:

$$
R=A\left(\frac{1}{d}\right)^N
\tag{4}
$$

where the $+~C$ term has been omitted because, as observed earlier, it's fit value was zero within experimental uncertainty.  If we take the logarithm of both sides of Eq. (4), we get:

\begin{align}
\log(R) &= \log\left[A\left(\frac{1}{d}\right)^N\right]\\
& = \log(A) + \log\left[\left(\frac{1}{d}\right)^N\right]\\
\therefore \log(R) & = N\log\left(\frac{1}{d}\right) + \log(A)
\tag{5}
\end{align}

Notice that Eq. (5) is of the form of a straight line, $y = mx + b$, where:

\begin{align}
y & = \log(R)\\
m &= N\\
x & = \log\left(\frac{1}{d}\right)\\
b &= \log(A)
\end{align}

In other words, when data that obey a power law ($p = Aq^N$) are plotted as $\log(p)$ versus $\log(q)$, a linear trend will be observed and the slope will give the value of the power $N$ and the intercept will be equal to $\log(A)$.

### Summary

So-called log-log plots can be useful in two distinct situations:
 - When the data span many orders of magnitude (along either the horizontal or vertical axes, or both), log-log plots allow us to see details in the data that can be hidden in linear plots.
 - When the measured data are expected to follow a power law ($p = Aq^N$), log-log plots will produce a linear dataset whose slope is equal to the power $N$ and whose intercept is equal to $\log(A)$.

The hydraulic resistance data collected in Lab #4 span over three orders of magnitude and, from **Question 1.3**, we know that $R_\mathrm{H}$ versus $1/d$ follow a power law.  Therefore, the $R_\mathrm{H}$ versus $1/d$ data from Lab #4 are ideally suited for an analysis of the form $\log(R_\mathrm{H})$ versus $\log(1/d)$.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.2</span>** **<span style="color:red">(1 mark)</span>** 

Use the 'Scatter()' function from PHYS121.py to plot $\log(R_\mathrm{H})$ as a function of $\log(1/d)$.  

The calculation of the uncertainty in $\log(R_\mathrm{H})$ is not easily done using the propogation of uncertainty rules that you learned in PHYS 111.  For this reason, the calculation of the uncertainty in $\log(R_\mathrm{H})$ has been done for you.  (If you're interested, the uncertainty calculation makes use of the facts that: (1) $\log(y) = \ln(y)/\ln(10)$ and (2) if $z=\ln(y)$, then $\Delta z = \Delta y/y$.)

**REMEMBER** that in Python $\log(x)$ must be entered as:
```python
np.log10(x)
```

Fill in the missing data needed when calling the 'Scatter()' function.  Remember that the scatter function accepts arrays or lists as inputs - to call a column from a DataFrame as a list, use the syntax
```python
df["Column Name"]
```
where 'df' is the name of the DataFrame in question. For additional information about the 'Scatter()' function, see **Appendix A** near the end of this notebook.   
 
 *** Please do not change any of the variable names in the cell below ***

```python
PHYS121.Scatter(xData = ..., 
                yData = ..., 
                yErrors = Lab4Data["RH_uncertainty"]/(Lab4Data["RH"]*np.log(10)), 
                xlabel = "log(1/d)", 
                ylabel = "log(R_H)"
                );
```

<!-- END QUESTION -->

Compared to the scatter plot of **Question 1.2**, the plot above should allow you to better see the details of the data at small values of $1/d$. 

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.3</span>** **<span style="color:red">(1 mark)</span>** 

Use the 'LinearFit()' function from PHYS121.py to plot $\log(R_\mathrm{H})$ versus $\log(1/d)$ and fit the data to a straight line.  

Again, the calculation of the uncertainty in $\log(R_\mathrm{H})$ has already been done for you.

**REMEMBER** that in Python $\log(x)$ must be entered as:
```python
np.log10(x)
```

Fill in the missing data needed when calling the 'LinearFit()' function.  Remember that the 'LinearFit()' function accepts arrays or lists as inputs - to call a column from a DataFrame as a list, use the syntax:
```python
df["Column Name"]
```
where 'df' is the name of the DataFrame in question. For additional information about the 'LinearFit()' function, see **Appendix C** near the end of this notebook.   
 
 *** Please do not change any of the variable names in the cell below ***

```python
PHYS121.LinearFit(xData = ..., 
                yData = ..., 
                yErrors = Lab4Data["RH_uncertainty"]/(Lab4Data["RH"]*np.log(10)), 
                xlabel = "log(1/d)", 
                ylabel = "log(R_H)"
                );
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.4</span>** **<span style="color:red">(1 mark)</span>**   

From the fit above, what is the value of the power $N$ that determines how the hydraulic resistance depends of the inverse of the inner diameter of the Teflon tubes: $R_\mathrm{H} = A(1/d)^N$?  How does this value compare to the value obtained in **Question 1.3**?  If there are differences, think about and describe what might be causing the difference in the two values of $N$. 

***
**<span style="color:blue">Answer 2.4:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Hydraulic Capacitance (60 minutes)</span></div>

Recall from Lab #4 the analogy between electric and hydraulic circuits as summarized in the table below:

| <div style="width:150px">electric</div> | <div style="width:50px">&nbsp;</div> | <div style="width:190px">hydraulic</div> |
| --: | :-: | :-- |
charge: $q\rm\ [C]$ | $\Longleftrightarrow$ | volume: $v\rm\ [m^3]$
current: $ I = dq/dt\rm\ [A = C/s]$ | $\Longleftrightarrow$ | volume flow rate: $Q = dv/dt\rm\ [m^3/s]$
voltage: $\Delta V\rm\ [V]$ | $\Longleftrightarrow$ | pressure: $\Delta P\rm\ [Pa = N/m^2]$
resistance: $R=\Delta V/I\rm\ [\Omega = V/A]$ | $\Longleftrightarrow$ | hydraulic resistance: $R_\mathrm{H}=\Delta P/Q\rm\ [kg/m^4 s]$

We now want to build on this analogy by including capacitance.  In electric circuits, capacitance is defined by:

$$
C = \frac{q}{\Delta V}
\tag{6}
$$

where $q$ is the charge stored by the capacitor and the $\Delta V$ is the voltage difference across the capacitor.

For the hydraulic analogue, our capacitor $C_\mathrm{H}$ will store a volume $v$ of water and there will be a pressure difference $\Delta P$ across the capacitor.  Therefore, hydraulic capacitance is defined by:

$$
C_\mathrm{H} = \frac{v}{\Delta P}
\tag{7}
$$

An analysis of the units reveals that hydraulic capacitance is measured in units of $[C_\mathrm{H}] =\rm\ m^4 s^2/kg$.

An electric capacitor can be charged using a battery, resistor, and switch as shown in Fig. 4(a) below.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-5/images/charging.png" alt="image info" style="width: 68%; height: auto;" /><br>
    <b>Fig. 4: (a) Charging an electric capacitor using a battery, resistor, and switch. (b) The analogous charging of a hydraulic capacitor using a source of constant pressure (hydraulic battery), a narrow tube (hydraulic resistance), and valve.</b></center>
</p>

A schematic diagram of the hydraulic analogue that we'll use in today's lab is shown in Fig. 4(b).  The same water vessel will be used to provide an approximately constant pressure difference and you'll use the same Teflon tubes that used in Lab #4 as hydraulic resistors.  The station you're working at should have the same size tubes that you used in the last lab.  Confirm that the diameters of tubes match the diameters that you used in Lab #4.  If they don't match, let your TA know.  For our hydraulic capacitor, we'll use a glass tube.  The inner diameter of the glass tube is much larger than the inner diameters of Teflon tubes, so it doesn't offer much resistance to the flow of water. We'll assume that the resistance of our hydraulic capacitors are negligible.  Figure 5 shows a schematic of a hydraulic capacitor "charged" with a volume $v$ of water.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-5/images/CH.png" alt="image info" style="width: 10%; height: auto;" /><br>
    <b>Fig. 5: A hydraulic capacitor of cross-sectional area $A$ filled with water to a height of $\Delta h$.</b></center>
</p>

If the height of the water is $\Delta h$ and the hydraulic capacitor has a cross-sectional area $A=\pi r^2$, where $r$ is the inner radius of the tube, then the volume of water stored is $v = \pi r^2\Delta h$ and the pressure difference between the top and bottom surfaces of the water is $\Delta P = \rho g \Delta h$, where $\rho = 998\rm\ kg/m^3$ is the density of water and $g = 9.81\rm\ m/s^2$ is the gravitational acceleration.  Subbing these expression into the definition of hydraulic capacitance reveals that:

$$
C_\mathrm{H} = \frac{v}{\Delta P} = \frac{\pi r^2 \Delta h}{\rho g \Delta h} = \frac{\pi r^2}{\rho g}
\tag{8}
$$

which allows us to calculate the hydraulic capacitance of the glass tube in terms of constants that we know.

In class we showed that, for a capacitor $C$ charging through a resistor $R$, the stored charge evolves with time as:

$$
q(t) = q_\infty\left(1 - e^{-t/\tau}\right)
\tag{9}
$$

where $q_\infty$ is the equilibrium charge on the capacitor after the switch has been closed for a long time (i.e. $t\to\infty$) and $\tau = RC$ is a time constant which determines how quickly the capacitor charges/discharges.  The capacitor charges to $95\%$ of its full charge ($q_\infty$) in a time $t = 3\tau$.  

The volume of water stored in a hydraulic capacitor follows the same trend:

$$
v(t) = v_\infty\left(1 - e^{-t/\tau_\mathrm{H}}\right)
$$

where the hydraulic time constant is given by $\tau_\mathrm{H} = R_\mathrm{H} C_\mathrm{H}$.  Because volume is directly proportional to the height of the water column in the hydraulic capacitor, it is also true that:

$$
h(t) = h_\infty\left(1 - e^{-t/\tau_\mathrm{H}}\right)
\tag{10}
$$


For this part of the lab:
 - You will first set up the hydraulic circuit shown in Fig. 4(b).  
 - With the hydraulic capacitor initially uncharged (empty), you will open the valve and record the height $h(t)$ of the water filling the hydraulic capacitor as a function of time.  
 - You will then fit the data to Eq. (9) in order to determine the time constant $\tau_\mathrm{H} = R_\mathrm{H} C_\mathrm{H}$ of the hydraulic circuit.
 - Using the known value of $R_\mathrm{H}$ from Lab #4, you will determine the hydraulic capacitance $C_\mathrm{H}$ from the measured time constant.
 - Finally, you will compare the experimentally-determined value of $C_\mathrm{H}$ to the expected value given by Eq. (8).
 - You will do all of these steps for both of the hydraulic resistors at your station.
 
Here, once again, is a summary of the analogy between electric and hydraulic circuits.  The last three rows of the table below have been added to include capacitance and the time constant.   In the second-to-last row, $d$ is the separation distance between a pair of parallel plates of area $A$. 

| <div style="width:150px">electric</div> | <div style="width:50px">&nbsp;</div> | <div style="width:190px">hydraulic</div> |
| --: | :-: | :-- |
charge: $q\rm\ [C]$ | $\Longleftrightarrow$ | volume: $v\rm\ [m^3]$
current: $ I = dq/dt\rm\ [A = C/s]$ | $\Longleftrightarrow$ | volume flow rate: $Q = dv/dt\rm\ [m^3/s]$
voltage: $\Delta V\rm\ [V]$ | $\Longleftrightarrow$ | pressure: $\Delta P\rm\ [Pa = N/m^2]$
resistance: $R=\Delta V/I\rm\ [\Omega = V/A]$ | $\Longleftrightarrow$ | hydraulic resistance: $R_\mathrm{H}=\Delta P/Q\rm\ [kg/m^4 s]$
capacitance: $C=Q/\Delta V\rm\ [F = C/V]$ | $\Longleftrightarrow$ | hydraulic capacitance: $C_\mathrm{H}=v/\Delta P\rm\ [m^4s^2/kg]$
parallel-plate capacitor: $C=\varepsilon_0A/d\rm\ [F = C/V]$ | $\Longleftrightarrow$ | hydraulic tube capacitor: $C_\mathrm{H}=\pi r^2/(\rho g)\rm\ [m^4s^2/kg]$
time constant: $\tau = RC\rm\ [s]$ | $\Longleftrightarrow$ | hydraulic time constant: $\tau_\mathrm{H} = R_\mathrm{H}C_\mathrm{H}\rm\ [s]$

## <div style="background-color:rgba(0, 0, 255, 0.05); vertical-align: middle; padding:15px 0; margin-top:5px">TIPS...<div>

### Draining the Hydraulic Capacitor

When changing the Teflon tube, or when repeating a trial, you'll need to drain the water from the hydraulic capacitor/glass tube.  Follow these steps to drain the water without making a mess.

 - First, close the valve of the large water vessel.
 - Next, take the glass beaker at your station and hold it beneath the spout of the large water vessel and pull out the orange rubber stopper.
 - Keep the end of the Teflon tube/orange stopper above the beaker and lower them together so that the water quickly drains from the glass tube and into the beaker.
 - Next, hold the beaker beneanth the glass tube and pull the rubber stopper from the bottom of the glass tube.  Use the beaker to catch any remaining water that falls from the glass tube.
 - Use a towel to soak up any residual water clinging to the bottom of the glass tube.
 - Dump the water collected in the beaker back into the top of the water vessel.
 - Finally, replace the Teflon tube (or insert a new one) and start your next trial.
 
 ### Measuring the Water Height vs Time
 
For the largest Teflon tubes, the hydraulic capacitor will fill quickly (at least initially).  You and your partner can devise a strategy to collect high-quality data.  You may find it necessary to revise your strategy and repeat your data collection several times.  That's fine!  Getting good data on the first attempt has no real value.  What's most important is to be patient and persistent to ensure that you get a good dataset for the final trial!
 
One strategy that you can consider trying is using your phone to record the water level as it rises.  This strategy will work best if you keep the height of your camera lens at approximately the same height as the top surface of the water filling the glass tube.  That is, rasie your phone at the same rate that the water level rises.  You can then play the video back at record heights and times.

In case its helpful, here are the links to the YouTube videos from Lab #4 that play beeps are regular time intervals.  Note, howver, for this lab it is not important for the measurements of the water height to be spaced equally in time.

| interval | link |
--: | :--
$5\rm\ s$ | https://www.youtube.com/watch?v=NMIZQcebfKI
$10\rm\ s$ | https://www.youtube.com/watch?v=Xjv2-F6Quls
$15\rm\ s$ | https://www.youtube.com/watch?v=LaUe-Hf0uSM
$30\rm\ s$ | https://www.youtube.com/watch?v=2AFro8meJGI
$60\rm\ s$ | https://www.youtube.com/watch?v=-22N-qs5PHE
$120\rm\ s$ | https://www.youtube.com/watch?v=2dAorgAB0I4


<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.1</span>** **<span style="color:red">(1 mark)</span>**  

(i) For the first Teflon tube, collect $h(t)$ versus $t$ data (in *centimeters* and *seconds*, respectively) and enter the measurements into the "tube1" DataFrame below.  It is best if you wait long enough for the hydraulic capacitor to fully charge (i.e. until the water level in the glass tube is no longer increasing).  You'll have to determine yourself how long is long enough.  Collect at least $15$ data points.

(ii) Enter your Teflon tube's diameter (in *milimeters*) in the "tube1_diameter" variable below.

(iii) Return to the [Google Sheets document](https://docs.google.com/spreadsheets/d/1BCeJboA1V6BY9LYW8Yyb8QwEOUXmWDtQglrZTukf02c/edit?usp=sharing) used at the end of Lab #4.  Use this sheet to enter in the hydraulic resistance (and its uncertainty) that you found for this diameter of Teflon tube. 

*** Please do not change any of the variable names in the cell below ***

```python
# Tube 1 DataFrame
tube1_diameter = ... # mm
RH1 = ...
dRH1 = ...
tube1 = pd.DataFrame({"time" : [...],
                      "height" : [...]
                      })
print(f"Results for tube with diameter {tube1_diameter} mm:")
tube1
```

```python
PHYS121.graderCheck([tube1_diameter, RH1, dRH1, tube1.to_numpy()], ['tube1_diameter', 'RH1', 'dRH1', 'tube1'], grader.check('q3.1'))
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.2</span>** **<span style="color:red">(1 mark)</span>**  

Create a scatter plot of the water height in the hydraulic capacitor versus time along with appropriate axis labels, with units, for tube 1.

For additional information about the 'Scatter' function, refer to **Appendix A** near the end of this notebook.

 *** Please do not change any of the variable names in the cell below ***

```python
fig5 = PHYS121.Scatter( xData = ...,
                        yData = ..., 
                        xlabel = "Time", 
                        xUnits = "s", 
                        ylabel = "Height", 
                        yUnits = "cm") 
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.3</span>** **<span style="color:red">(1 mark)</span>**  

Fit the data for tube 1 to a function of the form:
$$
h(t)=h_\infty\left(1 - e^{-t/\tau_\mathrm{H}}\right)
$$

which is a copy of Eq. (9) above.  The 'Charging()' function from PHYS121.py will be used to plot $h(t)$ as a function of $t$ and fit the data.  

The 'Charging()' function is called as follows:
```python
import PHYS121
A, tau, dA, dtau, fig = PHYS121.Charging(xData = x, yData = y, yErrors = erry, xlabel = 'x', ylabel = 'y', xUnits ='s', yUnits = 'm')
```

For the fits in this part of the lab, we will neglect the uncertainty in the $h(t)$ measurements.  For additional information about the 'Charging()' function, refer to **Appendix D** near the end of this notebook.

 *** Please do not change any of the variable names in the cell below ***

```python
PHYS121.Charging( xData = ...,
                yData = ..., 
                xlabel = "Time", 
                xUnits = "s", 
                ylabel = "Height", 
                yUnits = "cm"); 
```

<!-- END QUESTION -->

***
**<span style="color:blue">Question 3.4</span>** **<span style="color:red">(1 mark)</span>**  

Based on the fit result from **<span style="color:blue">Question 3.3</span>**, enter the experimentally-determined value of the time constant $\tau_1$ and its uncertaintly $\Delta\tau_1$.

***
**<span style="color:blue">Answer 3.4:</span>**

In the cell below, replace the instances of ... with the appropriate numerical value for $\tau_1$ and $\Delta \tau_1$.

 *** Please do not change any of the variable names in the cell below ***

```python
tau1 = ...
dtau1 = ...
print(f"Results for tube with diameter {tube1_diameter} mm:")
print(f"tau1 = ({tau1:.5g} +/- {dtau1:.5g}) s")
```

```python
PHYS121.graderCheck([tau1, dtau1], ['tau1', 'dtau1'], grader.check('q3.4'))
```

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.5</span>** **<span style="color:red">(1 mark)</span>**  

Next, for the first Teflon tube, use the experimental values of $\tau_\mathrm{H}\pm \Delta \tau_\mathrm{H}$ and $R_\mathrm{H}\pm\Delta R_\mathrm{H}$ to determine the hydraulic capacitance $C_\mathrm{H}\pm\Delta C_\mathrm{H}$.

Note that the numerical values for $\tau_\mathrm{H}\pm\Delta\tau_\mathrm{H}$ and $R_\mathrm{H}\pm\Delta R_\mathrm{H}$ have a already been stored in variables called: 
```python
tau1
dtau1
RH1
dRH1
```

***
**<span style="color:blue">Answer 3.5:</span>**

Make use of these variables.  It is not necessary to type in the numerical values.  In the cell below, replace the instances of ... with the appropriate expression to calculate $C_\mathrm{H}$ and $\Delta C_\mathrm{H}$ for the data collected using the first Teflon tube. 

 *** Please do not change any of the variable names in the cell below ***

```python
CH1 = ...
dCH1 = ...
print(f"Results for tube with diameter {tube1_diameter} mm:")
print(f"CH2 = ({CH1:.3g} +/- {dCH1:.3g}) m^4s^2/kg")
```

```python
PHYS121.graderCheck([CH1, dCH1], ['CH1', 'dCH1'], grader.check('q3.5'))
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.6</span>** **<span style="color:red">(1 mark)</span>**  

(i) For the second Teflon tube, collect $h(t)$ versus $t$ data (in *centimeters* and *seconds*, respectively) and enter the measurements into the "tube2" DataFrame below.  It is best if you wait long enough for the hydraulic capacitor to fully charge (i.e. until the water level in the glass tube is no longer increasing).  You'll have to determine yourself how long is long enough.  Collect at least $15$ data points.

(ii) Enter your Teflon tube's diameter (in *milimeters*) in the "tube2_diameter" variable below.

(iii) Return to the [Google Sheets document](https://docs.google.com/spreadsheets/d/1BCeJboA1V6BY9LYW8Yyb8QwEOUXmWDtQglrZTukf02c/edit?usp=sharing) used at the end of Lab #4.  Use this sheet to enter in the hydraulic resistance (and its uncertainty) that you found for this diameter of Teflon tube. 

*** Please do not change any of the variable names in the cell below ***

```python
# Tube 2 DataFrame
tube2_diameter = ... # mm
RH2 = ...
dRH2 = ...
tube2 = pd.DataFrame({"time" : [...],
                      "height" : [...]
                      })
print(f"Results for tube with diameter {tube2_diameter} mm:")
tube2
```

```python
PHYS121.graderCheck([tube2_diameter, RH2, dRH2, tube2.to_numpy()], ['tube2_diameter', 'RH2', 'dRH2', 'tube2'], grader.check('q3.6'))
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.7</span>** **<span style="color:red">(1 mark)</span>**  

Create a scatter plot of the water height in the hydraulic capacitor versus time along with appropriate axis labels, with units, for tube 2.

For additional information about the 'Scatter' function, refer to **Appendix A** near the end of this notebook.

 *** Please do not change any of the variable names in the cell below ***

```python
fig7 = PHYS121.Scatter( xData = ...,
                        yData = ..., 
                        xlabel = "Time", 
                        xUnits = "s", 
                        ylabel = "Height", 
                        yUnits = "cm") 
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.8</span>** **<span style="color:red">(1 mark)</span>**  

Fit the data for tube 2 to a function of the form:
$$
h(t)=h_\infty\left(1 - e^{-t/\tau_\mathrm{H}}\right)
$$

which is a copy of Eq. (9) above.  The 'Charging()' function from PHYS121.py will be used to plot $h(t)$ as a function of $t$ and fit the data.  

The 'Charging()' function is called as follows:
```python
import PHYS121
A, tau, dA, dtau, fig = PHYS121.Charging(xData = x, yData = y, yErrors = erry, xlabel = 'x', ylabel = 'y', xUnits ='s', yUnits = 'm')
```

For the fits in this part of the lab, we will neglect the uncertainty in the $h(t)$ measurements.  For additional information about the 'Charging()' function, refer to **Appendix D** near the end of this notebook.

 *** Please do not change any of the variable names in the cell below ***

```python
PHYS121.Charging( xData = ...,
                yData = ..., 
                xlabel = "Time", 
                xUnits = "s", 
                ylabel = "Height", 
                yUnits = "cm");
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.9</span>** **<span style="color:red">(1 mark)</span>**  

Based on the fit result from **<span style="color:blue">Question 3.8</span>**, enter the experimentally-determined value of the time constant $\tau_2$ and its uncertaintly $\Delta\tau_2$.

***
**<span style="color:blue">Answer 3.9:</span>**

In the cell below, replace the instances of ... with the appropriate numerical value for $\tau_2$ and $\Delta \tau_2$.

 *** Please do not change any of the variable names in the cell below ***

```python
tau2 = ...
dtau2 = ...
print(f"Results for tube with diameter {tube2_diameter} mm:")
print(f"tau2 = ({tau2:.5g} +/- {dtau2:.5g}) s")
```

```python
PHYS121.graderCheck([tau2, dtau2], ['tau2', 'dtau2'], grader.check('q3.9'))
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.10</span>** **<span style="color:red">(1 mark)</span>**  

Next, for the second Teflon tube, use the experimental values of $\tau_\mathrm{H}\pm \Delta \tau_\mathrm{H}$ and $R_\mathrm{H}\pm\Delta R_\mathrm{H}$ to determine the hydraulic capacitance $C_\mathrm{H}\pm\Delta C_\mathrm{H}$.

Note that the numerical values for $\tau_\mathrm{H}\pm\Delta\tau_\mathrm{H}$ and $R_\mathrm{H}\pm\Delta R_\mathrm{H}$ have a already been stored in variables called: 
```python
tau2
dtau2
RH2
dRH2
```

***
**<span style="color:blue">Answer 3.10:</span>**

Make use of these variables.  It is not necessary to type in the numerical values.  In the cell below, replace the instances of ... with the appropriate expression to calculate $C_\mathrm{H}$ and $\Delta C_\mathrm{H}$ for the data collected using the first Teflon tube. 

 *** Please do not change any of the variable names in the cell below ***

```python
CH2 = ...
dCH2 = ...
print(f"Results for tube with diameter {tube2_diameter} mm:")
print(f"CH2 = ({CH2:.2e} +/- {dCH2:.2e}) m^4s^2/kg")
```

```python
PHYS121.graderCheck([CH2, dCH2], ['CH2', 'dCH2'], grader.check('q3.10'))
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.11</span>** **<span style="color:red">(1 mark)</span>**  

Finally, use Eq. (8) to calculate the expected hydraulic capacitance of the glass tube.  To complete this calculation, you'll need to know that the inner radius of the glass tube is $r = 4.36\rm\ mm$.

***
**<span style="color:blue">Answer 3.11:</span>**

In the cell below, replace the ... with an expression that calculates the expected value $C_\mathrm{H}$.  Recall that $\pi$ can be accessed in Python using:
```python
np.pi
```
and powers are entered by repeating the multiplication sign as in:
```python
x**2
```

 *** Please do not change any of the variable names in the cell below ***

```python
r = 4.36e-3 # m
rho = 998 # kg/m^3
g = 9.81 # m/s^2
CH_expected = ...
print(f"The expected hydraulic capacitance is {CH_expected:.2e} m^4s^2/kg.")
```

```python
PHYS121.graderCheck([CH_expected], ['CH2_expected'], grader.check('q3.11'))
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.12</span>** **<span style="color:red">(1 mark)</span>**   

 - Compare the two experimental values for $C_\mathrm{H}\pm\Delta C_\mathrm{H}$.  Do the values agree within experimental uncertainty?  If not, discuss some possible reasons that could account for the differences.
 
 - Compare the two experimental values for $C_\mathrm{H}\pm\Delta C_\mathrm{H}$ to the expected value calculated in **Question 3.9**.  Do the measured values agree with the exepcted value within experimental uncertainty?  If not, discuss some possible reasons that could account for the differences.

***
**<span style="color:blue">Answer 3.12:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - Stored Energy (30 minutes)</span></div>

One way to express the energy stored by a charged capacitor is:

$$
U = \frac{1}{2}q\,\Delta V
\tag{11}
$$

In the hydraulic analogue, $q\to v$ and $\Delta V\to \Delta P$ such that:

$$
U_\mathrm{H} = \frac{1}{2}v\,\Delta P
\tag{12}
$$

For a hydraulic capactor charged to a height difference $\Delta h$, the pressure difference across the capacitor is given by $\Delta P = \rho g\,\Delta h$.  Subbing this result into Eq. (12) are reordering the factors gives:

$$
U_\mathrm{H} = \frac{1}{2}(\rho v) g\Delta h = mg\left(\frac{\Delta h}{2}\right)
\tag{13}
$$

where, in the last step, we have used the fact that $\rho v$ is equal to the mass of the stored water.

Recall that $mgy$ is the gravitational potential energy of a point mass $m$.  Notice also that $\Delta h/2$ is the average height of the column of water in the hydraulic capacitor. See Fig. 6(b) below.  These observations show that the energy stored by a hydraulic capacitor corresponds to the average *gravitational potential energy* of the stored water.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/Lab-5/images/storedEnergy.png" alt="image info" style="width: 80%; height: auto;" /><br>
    <b>Fig. 6: (a) An initially uncharged hydraulic capacitor. (b) A long time after openning the valve, the hydraulic capacitor fully charges.  A mass $m$ of water (shaded red) is effectively removed from the top of the water vessel (at a height $\Delta h$) and transferred into the the hydraulic capacitor.</b></center>
</p>

We can also calculate the energy *lost* by the hydraulic battery during the charging process.  Effectively, as shown in Fig. 6, a mass $m$ of water was transferred from the top of the water vessel into the hydraulic capacitor.  Since the transferred mass was all effectively at a height $\Delta h$, the energy lost by the hydraulic battery is:

$$
U_\mathrm{lost} = mg\Delta h
\tag{14}
$$

which is twice the energy stored by the capacitor!

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 4.1</span>** **<span style="color:red">(1 mark)</span>**   

The analysis above reveals that charging the hydraulic capacitor causes the hydraulic battery to loss twice as much energy ($mg\Delta h$) as the hydraulic capacitor gains ($mg\Delta h/2$).  

 - What is the resolution of this apparent discrepancy?  
 - Why does the battery loss twice the energy that the capacitor gains?  
 - Where has the other half of the energy gone?
 
Discuss some possible explanations with your partner and perhaps some of the neighbouring groups.  Enter your thoughts and ideas into the cell below. 

***
**<span style="color:blue">Answer 4.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 5 - Reflection and statement of contributions (20 minutes)</span></div>

Hopefully Labs #4 and 5 have shown that concepts from electric circuits can be applied to problems in fluid dynamics:
 - Voltage differences drive the flow of charge through resistors in electric circuits $\Longleftrightarrow$ Pressure differences drive the flow of fluid through pipes in hydraulic circuits
 - Electric capacitors store a quantity of charge and electrostatic potential energy $\Longleftrightarrow$ Hydraulic capacitors store a volume of fluid and gravitational potential energy
 
The analogue can be extended to other fields of physics.
 - Temperature difference drive the flow of heat through solid materials in thermal circuts.
 - Thermal capacitors store heat and thermal energy. 
 
However, it is important to also recognize that there are also limitations to the analogy.  In **Part 1** of today's lab, you likely found that, while electric resitance is proportional to $d^{-2}$, hydraulic resistance is proportional to $d^{-4}$.  This is an important difference, especially if you're designing a complex system of pumps and pipes.  

***
Executing the cell below will open a YouTube video that shows a large electric capacitor charging through a $3000\rm\ \Omega$ resistor.  The video is almost 27 minutes long.  You do **not** have to watch the whole thing, we're not going to ask you any questions about it.  Just use the scroll bar to move around to different point in the video.

There are two multimeters in the video.  The one on the left measures the constant voltage output by the power supply.  It acts as the battery and is analogous to the large water vessel you used the today's lab and in Lab #4.  The multimeter on the right monitors the voltage across the capacitor which is expected to evolve with time as:

$$
V(t) = V_\infty\left(1-e^{-t/\tau}\right)
$$

What you should notice from watching the video is that initially the capacitor voltage increases relatively quickly.  The charging rate then decreases as time goes on.  You observed the exact same phenomenon in today's lab when measuring the height of the water column in the hydraulic capacitor.

If you're looking for a minor challenge, try using the video and the fact that the capacitor is charging through a resistance of $R=3000\rm\ \Omega$ to determine the value of the large blue capacitor.  Note: you do **not** need to collect a lot of data or do a plot/fit to find the value of $C$.

```python
YouTubeVideo("8DtKaZ2XiHA", width=500)
```

Finally, executing the cell below will take you directly to the part of the video where the capacitor, after being fully charged, is quickly discharged.  Because the capacitor is large, it stores a lot of charge.  By rapidly discharging the capacitor, the stored energy can be reovered in a very short time interval.  In the video, the capacitor is first disconnected from the power supply and multimeters.  Then, after turning out the lights, a screw driver is used to discharge the capacitor by short-circuitng the two capacitor terminals.

For the hydraulic circuit, short circuiting the terminals of the capacitor is equivalent to pulling the stopper from the bottom of the glass tube after it has been fully charged.  The rush of water that comes out the bottom could be used to do useful work (such as spinning a turbine).

```python
YouTubeVideo("8DtKaZ2XiHA", width=500, start=1524)
```

***
Below, you will briefly summarize the contributions that you and your partner made when completely this lab.  Please ensure that both partners have a chance to engage in all types of lab activities (developing the experimental methods/data collection strategies, hands-on interaction with the apparatus/data collection, and data processing/analysis in the Jupyter notebook).  Note that statements of contributions are common in science.  The figure below shows one example taken from a 2021 paper published in the journal *Nature Physics*.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/contributions.jpg" alt="image info" style="width: 60%; height: auto;" /><br>
<b>Fig. 2: A sample statement of author contributions from a recent journal article.</b></center>
</p>


<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 5.1(a):</span>**  

Answer this question only if you worked with a partner to complete this lab.

In two or three sentences describe the specific contributions of **Partner A** to this lab.  For example, worked with the experimental apparatus to collect data for Part 1 of the lab, entered and processed the data in the Jupyter notebook for Part 2 of the lab, ...  

***
**<span style="color:blue">Answer 5.1(a):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  Please start be entering the name and student number of **Partner A**.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 5.1(b):</span>**  

Answer this question only if you worked with a partner to complete this lab.

In two or three sentences describe the specific contributions of **Partner B** to this lab.  

***
**<span style="color:blue">Answer 5.1(b):</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  Please start be entering the name and student number of **Partner B**.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 6 - Submission (5 minutes)</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 6.1:</span>**  

We welcome your feedback on the PHYS 121 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 121 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 6.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Canvas gradebook.
- **Do NOT change the number of the .zip file.**
- **Do NOT modify the contents of the .zip file.**

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(files = ['PHYS121_DataLogger.txt'])
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 7 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.  
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.  
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS121
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS121.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix B &ndash; The PowerLaw Function...</div>

The power law function fits data to a function of the form $y=A\,x^N + C$ and is called as follows:
```python
PowerLaw(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '')
```
The xData and yData lists are required.  The others are optional with default values set.  The function returns the following outputs:
```python
A, N, C, errA, errN, errC, fig
```
The first six outputs are numerical values and 'fig' is the formatted plot. 

If $y$-uncertainties are provided, the function will perform a weighted fit.  The $y$-uncertainties list must be the same length as $x$- and $y$-data lists.  If $y$-uncertainties are not provided, the fit will be unweighted.  The other optional arguments include $x$- and $y$-axis names and units for the $x$- and $y$-datasets.  These must be entered as strings (enclosed in quotations) and they are used for formatting the outputs of the function.

### Unweighted Fit Example Implmentation
The code block below shows the most basic use of 'PowerLaw' for an unweighted fit without any of the optional arguments.
```python
import PHYS121
x = [1, 2, 3, 4]
y = [1.1, 3.9, 8.7, 16.5]
A, N, C, dA, dN, dC, fig = PHYS121.PowerLaw(x, y)
```

### Weighted Fit Example Implmentation
This second blcok of code shows how to use 'PowerLaw' to do a weighted fit.  It also makes use of all of the other optional arguments.  
```python
import PHYS121
x = [1, 2, 3, 4]
y = [1.1, 3.9, 8.7, 16.5]
erry = [0.05, 0.12, 0.20, 0.5]
A, N, C, dA, dN, dC, fig = PHYS121.PowerLaw(x, y, erry, 'x' , 'y', 's', 'm')
```

The 'A, N, C, dA, dN, dC, fig = ...' syntax is used to store the function outputs in variables.
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix C &ndash; The LinearFit Function...</div>

The linear-fit function is called as follows:
```python
LinearFit(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False)
```
The xData and yData lists are required.  The others are optional with default values set.  The function returns the following outputs:
```python
Slope, Yintercept, errSlope, errYintercept, fig
```
The first four outputs are numerical values and 'fig' is the formatted plot. 

If $y$-uncertainities are provided, the function will perform a weighted fit.  The $y$-uncertainties list must be the same length as $x$- and $y$-data lists.  If $y$-uncertainties are not provided, the fit will be unweighted.  The other optional arguments include $x$- and $y$-axis names and units for the $x$- and $y$-datasets.  These must be entered as strings (enclosed in quotations) and they are used for formatting the outputs of the function.  If 'fill' is set to True, shading will be added around the best-fit line representing confindence intervals. 

### Unweighted Fit Example Implmentation
The code block below shows the most basic use of 'LinearFit' for an unweighted fit without any of the optional arguments.
```python
import PHYS121
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
m, b, dm, db, fig = PHYS121.LinearFit(V, I)
```

### Weighted Fit Example Implmentation
This second blcok of code shows how to use 'LinearFit' to do a weighted fit.  It also makes use of all of the other optional arguments.  
```python
import PHYS121
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
errI = [0.005, 0.012, 0.020, 0.025]
m, b, dm, db, fig = PHYS121.LinearFit(V, I, errI, 'voltage' , 'current', 'V', 'A', True)
```

The 'm, b, dm, db, fig = ...' syntax is used to store the function outputs in varibles.  If you're interested in details of the fitting process, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Weighted%20Linear%20Fit.html
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix D &ndash; The Charging Function...</div>

The Charging function is called as follows:
```python
Charging(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '')
```
The xData and yData lists are required.  The others are optional with default values set.  The function returns the following outputs:
```python
A, tau, errA, errtau, fig
```
The first four outputs are numerical values and 'fig' is the formatted plot. 

If $y$-uncertainities are provided, the function will perform a weighted fit.  The $y$-uncertainties list must be the same length as $x$- and $y$-data lists.  If $y$-uncertainties are not provided, the fit will be unweighted.  The other optional arguments include $x$- and $y$-axis names and units for the $x$- and $y$-datasets.  These must be entered as strings (enclosed in quotations) and they are used for formatting the outputs of the function.

### Unweighted Fit Example Implmentation
The code block below shows the most basic use of 'Charging' for an unweighted fit without any of the optional arguments.
```python
import PHYS121
t = [1, 2, 3, 4]
h = [0.095, 0.182, 0.255, 0.335]
A, tau, dA, dtau, fig = PHYS121.Charging(t, h)
```

### Weighted Fit Example Implmentation
This second blcok of code shows how to use 'Charging' to do a weighted fit.  It also makes use of all of the other optional arguments.  
```python
import PHYS121
t = [1, 2, 3, 4]
h = [0.095, 0.182, 0.255, 0.335]
errh = [0.005, 0.012, 0.020, 0.025]
A, tau, dA, dtau, fig = PHYS121.Charging(t, h, errh, 'time' , 'height', 's', 'cm')
```

The 'A, tau, dA, dtau, fig = ...' syntax is used to store the function outputs in varibles.  If you're interested in details of the nonlinear fitting process, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Weighted%20Nonlinear%20Fit.html
***

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: February 29, 2024

Tests for name_and_student_number (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...


Tests for q1.4 (points: None)
```
>>> a1_4.lower() in ["a", "b", "c", "d", "e", "f", "g", "h", "i"]
True
```
failure_message: Your answer must be a string of one of 'a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', or 'i'.


Tests for q2.1 (points: None)
```
>>> a2_1.lower() in ["a", "b", "c", "d", "e", "f"]
True
```
failure_message: Your answer must be a string of one of 'a', 'b', 'c', 'd', 'e', or 'f'.


Tests for q3.1 (points: None)
```
>>> tube1_diameter in [0.965, 1.397, 1.676, 1.930, 2.413]
True
```
failure_message: Enter your tube diameter in units of millimeters.  It must be one of 0.965, 1.397, 1.676, 1.930, or 2.413 mm
```
>>> isinstance(RH1, float)
True
```
failure_message: You must enter a decimal number for 'RH1'.
```
>>> isinstance(dRH1, float)
True
```
failure_message: You must enter a decimal number for 'dRH1'.
```
>>> len(tube1["time"]) >= 15 and len(tube1["height"]) >= 15 
True
```
failure_message: You must collect at least 15 data points.
```
>>> len(tube1["time"]) == len(tube1["height"]) 
True
```
failure_message: The number of elements in 'tube1["time"]' does not match the number of elements in 'tube1["height"]'.
```
>>> all(tube1["height"]) >= 0 and all(tube1["height"]) < 25
True
```
failure_message: The measured heights must be entered in units of centimeters and must all be positive.


Tests for q3.10 (points: None)
```
>>> isinstance(CH2, float)
True
```
failure_message: CH2 must evaluate to a decimal number.
```
>>> isinstance(dCH2, float)
True
```
failure_message: dCH2 must evaluate to a decimal number.
```
>>> np.isclose(CH2, tau2/RH2, rtol=1e-02)
True
```
failure_message: Try again.  Recall that tau = RH*CH.
```
>>> np.isclose(dCH2, CH2*(dtau2/tau2 + dRH2/RH2), rtol=1e-02)
True
```
failure_message: Try again.  Remember that fractional errors add for prducts and quotients.


Tests for q3.11 (points: None)
```
>>> np.isclose(CH_expected, np.pi*r**2/(rho*g), rtol=1e-02)
True
```
failure_message: Try again.  See Eq. (8).


Tests for q3.4 (points: None)
```
>>> isinstance(tau1, float)
True
```
failure_message: tau1 must be a decimal number.
```
>>> isinstance(dtau1, float)
True
```
failure_message: dtau1 must be a decimal number.
```
>>> def ChargingFcn(x, A, tau):
...     y = A*(1 - np.exp(-x/tau))
...     return y
>>> a_fit, cov = curve_fit(ChargingFcn, np.array(tube1["time"])/max(tube1["time"]), np.array(tube1["height"])/max(tube1["height"]))
>>> tau_fit = a_fit[1]*max(tube1["time"])
>>> np.isclose(tau1, tau_fit, rtol = 0.1)        
True
```
failure_message: Look at the fit results from Question 3.3 to find the correct value of tau.
```
>>> def ChargingFcn(x, A, tau):
...     y = A*(1 - np.exp(-x/tau))
...     return y
>>> a_fit, cov = curve_fit(ChargingFcn, np.array(tube1["time"])/max(tube1["time"]), np.array(tube1["height"])/max(tube1["height"]))
>>> errTau = np.sqrt(np.diag(cov))[1]*max(tube1["time"])
>>> np.isclose(dtau1, errTau, rtol = 0.1)      
True
```
failure_message: Look at the fit results from Question 3.3 to find the correct value of tau.


Tests for q3.5 (points: None)
```
>>> isinstance(CH1, float)
True
```
failure_message: CH1 must evaluate to a decimal number.
```
>>> isinstance(dCH1, float)
True
```
failure_message: dCH1 must evaluate to a decimal number.
```
>>> np.isclose(CH1, tau1/RH1, rtol=1e-02)
True
```
failure_message: Try again.  Recall that tau = RH*CH.
```
>>> np.isclose(dCH1, CH1*(dtau1/tau1 + dRH1/RH1), rtol=1e-02)
True
```
failure_message: Try again.  Remember that fractional errors add for prducts and quotients.


Tests for q3.6 (points: None)
```
>>> tube2_diameter in [0.965, 1.397, 1.676, 1.930, 2.413]
True
```
failure_message: Enter your tube diameter in units of millimeters.  It must be one of 0.965, 1.397, 1.676, 1.930, or 2.413 mm
```
>>> isinstance(RH2, float)
True
```
failure_message: You must enter a decimal number for 'RH2'.
```
>>> isinstance(dRH2, float)
True
```
failure_message: You must enter a decimal number for 'dRH2'.
```
>>> len(tube2["time"]) >= 15 and len(tube2["height"]) >= 15 
True
```
failure_message: You must collect at least 15 data points.
```
>>> len(tube2["time"]) == len(tube2["height"]) 
True
```
failure_message: The number of elements in 'tube2["time"]' does not match the number of elements in 'tube2["height"]'.
```
>>> all(tube2["height"]) >= 0 and all(tube2["height"]) < 25
True
```
failure_message: The measured heights must be entered in units of centimeters.


Tests for q3.9 (points: None)
```
>>> isinstance(tau2, float)
True
```
failure_message: tau2 must be a decimal number.
```
>>> isinstance(dtau2, float)
True
```
failure_message: dtau2 must be a decimal number.
```
>>> def ChargingFcn(x, A, tau):
...     y = A*(1 - np.exp(-x/tau))
...     return y
>>> a_fit, cov = curve_fit(ChargingFcn, np.array(tube2["time"])/max(tube2["time"]), np.array(tube2["height"])/max(tube2["height"]))
>>> tau_fit = a_fit[1]*max(tube2["time"])
>>> np.isclose(tau2, tau_fit, rtol = 0.1)    
True
```
failure_message: Look at the fit results from Question 3.8 to find the correct value of tau.
```
>>> def ChargingFcn(x, A, tau):
...     y = A*(1 - np.exp(-x/tau))
...     return y
>>> a_fit, cov = curve_fit(ChargingFcn, np.array(tube2["time"])/max(tube2["time"]), np.array(tube2["height"])/max(tube2["height"]))
>>> errTau = np.sqrt(np.diag(cov))[1]*max(tube2["time"])
>>> np.isclose(dtau2, errTau, rtol = 0.1)   
True
```
failure_message: Look at the fit results from Question 3.8 to find the correct value of tau.
